Glucose-stimulated KIF5B-driven microtubule sliding organizes microtubule networks in pancreatic β cells

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Specific comments:

(1) It is difficult to appreciate that there is a "peripheral sub-membrane microtubule array" as it is not well defined in the manuscript. This reviewer assumes that this is in the respective field clear. Yet, while it is appreciated that there is an increased amount of MTs close to the cytoplasmic membrane, the densities appear very variable along the membrane. Please provide a clear description in the Introduction what is meant with "peripheral sub-membrane microtubule array".

A definition has been added to the Introduction.

(2) The authors described a "consistent presence of a significant peripheral array in the C57BL/ 6J control mice, while the KO counterparts exhibited a partial loss of this peripheral bundle.

Specifically, the measured tubulin intensity at the cell periphery was significantly reduced in the KO mice compared to their wild-type counterparts". In vitro "control cells had convoluted nonradial MTs with a prominent sub-membrane array, typical for β cells (Fig. 2A), KIF5B-depleted cells featured extra-dense MTs in the cell center and sparse receding MTs at the periphery (Fig. 2B,C)". Please comment/discuss why in vivo there are no "extra-dense MTs in the cell center".

We now add a discussion of this point, which we believe could be a manifestation of 3D shape of a beta cell in tissue and/or compensatory mechanisms in organisms.

(3) Authors should include in the Discussion a paragraph discussing the fact that small changes in MT configuration can have strong effects.

A paragraph added to the discussion.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

(1) Figure 1: Even though the reviewer appreciates that minor changes of MT configuration have severe effects, still the overall effects appear minor (40 vs. <50% or 35% vs. around 28%). Notably, there are no statistically significant differences in the different groups in Fig. 1Suppl-Fig.1 D. This reviewer is not sure if the combination of many not significantly different data points can result in significant changes and this should be checked by a statistician. Authors should include in the Discussion a paragraph discussing the fact that small changes in MT configuration can have strong effects.

We have now added the requested paragraph to the discussion. Indeed, the differences are small, and the significance is only detected in a data set with a large sample size in Fig. 1J,K (combined data sets with smaller sizes from Fig. 1-Suppl-Fig.1 D), consistent with the fact that a larger sample size generally provides more power to detect an effect.

(2) Unfortunately, the authors cannot block kinesin-1 resulting in microtubule accumulation in the cell center and then release the block (best inhibiting microtubule formation), to show that the MTs accumulated in the cell center will be transported to the periphery.

This is indeed the case at the moment, yes.

Minor comments:

- Abstract: β-cells vs. β cells (and throughout the manuscript)

- Page 4: "MTOC, the Golgi, (Trogden et al. 2019), and"

- Page 5: "β-cell specific"

- MT-sliding vs. MT sliding

- Kinesin 1 vs. kinesin-1

- Page 6, line 1: "β cells. actively"

- Page 7: "a microtubule probe", should be "MT"

- Page 9: "1μm" vs. "1 μm"

- Page 10: "demonstrate a dramatic effect" recommended is: "demonstrate a marked effect"

- Page 13, line 1: dramatically vs. markedly

- Page 13, line 5: "50μm" vs. "50 μm" (in general, there should be a space between number and unit?)

- "37 degrees C" vs. "37{degree sign}C"

- Animal protocol number?

- "Mice were euthanized by isoflurane inhalation"? What concentration? How long? More details are needed (no cervical dislocation?).

- Antibodies: more identifiers are needed.

- Antibody information in Key reagents and under 5. Reagents and antibodies do not fit (1:500 and 1:1000).

Thank you, we corrected all relevant information now.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: In pancreatic islet β cells, molecular motors use cytoskeletal polymers microtubules as tracks for intracellular transport of insulin secretory granules. β-cell microtubule network has a complex architecture and is non-directional, which provide insulin granules at the cell periphery for rapid secretion response, yet to avoid over-secretion and subsequent hypoglycemia. We have previously characterized a peripheral sub-membrane microtubule array, which is critical for withdrawal of excessive insulin granules from the secretion sites. Microtubules in β cells originate at the Golgi in the cell interior, and how the peripheral array is formed is unknown. Using real-time imaging and photo-kinetics approaches in clonal mouse pancreatic β cells MIN6, we now demonstrate that kinesin KIF5B, a motor protein with a capacity to transport microtubules as cargos, slides existing microtubules to the cell periphery and aligns them to each other along the plasma membrane. Moreover, like many physiological β-cell features, microtubule sliding is facilitated by a high glucose stimulus. These new data, together with our previous report that in high glucose sub-membrane MT array is destabilized to allow for robust secretion, indicate that MT sliding is another integral part of glucose-triggered microtubule remodeling, likely replacing destabilized peripheral microtubules to prevent their loss over time and β-cell malfunction.

Revised version: In pancreatic islet β cells, molecular motors use cytoskeletal polymers microtubules as tracks for intracellular transport of insulin secretory granules. The β-cell microtubule network has a complex architecture and is non-directional, which provides insulin granules at the cell periphery for rapid secretion response, yet to avoid over-secretion and subsequent hypoglycemia. We have previously characterized a peripheral sub-membrane microtubule array, which is critical for the withdrawal of excessive insulin granules from the secretion sites. Microtubules in β cells originate at the Golgi in the cell interior, and how the peripheral array is formed is unknown. Using real-time imaging and photo-kinetics approaches in clonal mouse pancreatic β cells MIN6, we now demonstrate that kinesin KIF5B, a motor protein with a capacity to transport microtubules as cargos, slides existing microtubules to the cell periphery and aligns them to each other along the plasma membrane. Moreover, like many physiological β-cell features, microtubule sliding is facilitated by a high glucose stimulus. These new data, together with our previous report that in high glucose sub-membrane MT array is destabilized to allow for robust secretion, indicate that MT sliding is another integral part of glucose-triggered microtubule remodeling, likely replacing destabilized peripheral microtubules to prevent their loss over time and β-cell malfunction.

Introduction

Earlier version: The precise level of glucose-stimulated insulin secretion (GSIS) from pancreatic β cells is crucial for glucose homeostasis. On one hand, insufficient insulin secretion decreases glucose uptake by peripheral tissues, leading to diabetes. On the other hand, excessive secretion causes glucose depletion from the bloodstream and hypoglycemia. Not surprisingly, multiple levels of cellular regulation control the amount of insulin secretory granules (IGs) released on every stimulus. One level of this control is facilitated by microtubules (MTs), intracellular polymers which serve as tracks for intracellular transport of IGs and define how many IGs are positioned at the secretion sites (Desai and Mitchison 1997, Varadi et al. 2002, Heaslip et al. 2014). Thus, the architecture and dynamics of the MT network is instrumental for secretion control.

Revised version: The precise level of glucose-stimulated insulin secretion (GSIS) from pancreatic β cells is crucial for glucose homeostasis. On one hand, insufficient insulin secretion decreases glucose uptake by peripheral tissues, leading to diabetes. On the other hand, excessive secretion causes glucose depletion from the bloodstream and hypoglycemia. Not surprisingly, multiple levels of cellular regulation control the amount of insulin secretory granules (IGs) released on every stimulus. One level of this control is facilitated by microtubules (MTs), intracellular polymers that serve as tracks for intracellular transport of IGs and define how many IGs are positioned at the secretion sites (Desai and Mitchison 1997, Varadi et al. 2002, Heaslip et al. 2014). Thus, the architecture and dynamics of the MT network is instrumental for secretion control.

Earlier version: The findings in the last decade have uncovered that organization and regulation of the β cell MT network are quite unusual. As in several other eukaryotic cells, MTs in β cells are nucleated at MT-organizing centers (MTOCs) in the cell interior, partially at the centrosome and to a large extent at the Golgi membranes (Zhu et al. 2015, Trogden et al. 2019). Conventionally, this should be followed by MT plus-end polymerization toward the cell periphery and result in a radial MT array with high MT density in the center rather than in the periphery. However, the β cell lacks such MT polarity, well-characterized in mesenchymal cells in culture (Bracey et al. 2022). Instead, interior β-cell MTs are twisted and interlocked (Varadi et al. 2003, Zhu et al. 2015), while peripheral MTs are arranged in a prominent array of co-aligned MTs underlying cell membrane (Bracey et al. 2020).

Revised version: The findings in the last decade have uncovered that the organization and regulation of the β-cell MT network are quite unusual. As in several other eukaryotic cells, MTs in β cells are nucleated at MT-organizing centers (MTOCs) in the cell interior, partially at the centrosome and to a large extent at the Golgi membranes (Zhu et al. 2015, Trogden et al. 2019). Conventionally, this should be followed by MT plus-end polymerization toward the cell periphery, resulting in a radial MT array with high MT density in the center rather than in the periphery. However, the β cell lacks such MT polarity, well-characterized in mesenchymal cells in culture (Bracey et al. 2022). Instead, interior β-cell MTs are twisted and interlocked (Varadi et al. 2003, Zhu et al. 2015), while peripheral MTs are arranged in a prominent array of co-aligned and often bundled MTs underlying the cell membrane, hereafter called the sub-membrane MT array or the peripheral MT array (Bracey et al. 2020).

Earlier version: Like other features on β cell physiology, MTs in β cells are regulated by glucose metabolism. As protein polymers, MTs are capable of regulated polymerization and depolymerization, which allows for fine tuning of their organization. Under basal glucose conditions, β cell MTs are stable and undergo little turnover (Zhu et al. 2015, Ho et al. 2020). Sub-membrane MT arrays are especially highly stabilized by a MT-associated protein (MAP) tau, well-known for its neuronal functions (Ho et al. 2020). Upon a high glucose stimulus, tau is phosphorylated and submembrane MTs undergoes destabilization (Ho et al. 2020) and fragmentation, possibly via a MT-severing activity (Muller et al. 2021). MT destabilization is accompanied by facilitated formation of new MTs at the main β cell MTOC, the Golgi, (Trogden et al. 2019), and increased MT polymerization (Heaslip et al. 2014), which are thought to replace destabilized MTs and restore the MT network.

Revised version: Like other features on β-cell physiology, MTs in β cells are regulated by glucose metabolism. As protein polymers, MTs are capable of regulated polymerization and depolymerization, which allows for fine-tuning of their organization. Under basal glucose conditions, β-cell MTs are stable and undergo little turnover (Zhu et al. 2015, Ho et al. 2020). Sub-membrane MT arrays are especially highly stabilized by a MT-associated protein (MAP) tau, well-known for its neuronal functions (Ho et al. 2020). Upon a high glucose stimulus, tau is phosphorylated and sub-membrane MTs undergoes destabilization (Ho et al. 2020) and fragmentation, possibly via a MT-severing activity (Muller et al. 2021). MT destabilization is accompanied by facilitated formation of new MTs at the main β-cell MTOC, the Golgi (Trogden et al. 2019), and increased MT polymerization (Heaslip et al. 2014), which are thought to replace destabilized MTs and restore the MT network.

Earlier version: Such complex, glucose-dependent organization of MT networks in β cells is thought to underly a multi-faceted involvement of MT transport in insulin secretion regulation: MTs regulate the availability of IGs for secretion in both positive and negative fashion. Firstly, Golgi-derived MTs are necessary for efficient IG generation at the trans Golgi network (TGN) (Trogden et al. 2019). Secondly, interior MTs are to a large extent responsible for IG transportation throughout the cell, which occurs predominantly in non-directional, diffusion-like manner (Tabei et al. 2013, Zhu et al. 2015), likely due to the interlocked configuration of interior MT network. In addition, directional MT-dependent runs of secretion-competent granules toward cell periphery have been described (Hoboth et al. 2015, Muller et al. 2021). These processes contribute to the positive regulation of GSIS by preparing and distributing IG. At the same time, sub-membrane MT arrays serve for withdrawal of excessive peripheral IGs from the secretion sites, which prevents IG docking and acute over-secretion upon a given stimulus (Zhu et al. 2015, Bracey et al. 2020, Hu et al. 2021). The latter process provides a negative MT regulation of secretion.

Revised version: Such complex, glucose-dependent organization of MT networks in β cells is thought to underly a multi-faceted involvement of MT transport in insulin secretion regulation: MTs regulate the availability of IGs for secretion in both positive and negative fashion. Firstly, Golgi-derived MTs are necessary for efficient IG generation at the trans Golgi network (TGN) (Trogden et al. 2019). Secondly, interior MTs are to a large extent responsible for IG transportation throughout the cell, which occurs predominantly in the non-directional, diffusion-like manner (Tabei et al. 2013, Zhu et al. 2015), likely due to the interlocked configuration of the interior MT network. In addition, directional MT-dependent runs of secretion-competent granules toward the cell periphery have been described (Hoboth et al. 2015, Muller et al. 2021). These processes contribute to the positive regulation of GSIS by preparing and distributing IG. At the same time, sub-membrane MT arrays serve for withdrawal of excessive peripheral IGs from the secretion sites, which prevents IG docking and acute over-secretion upon a given stimulus (Zhu et al. 2015, Bracey et al. 2020, Hu et al. 2021). The latter process provides a negative MT regulation of secretion.

Earlier version: During GSIS, the above mentioned glucose-dependent destabilization of the submembrane MT array must reduce IG withdrawal and downplay the negative MT-dependent regulation of secretion(Bracey et al. 2020). At the same time, new MTs polymerizing off the Golgi facilitate IG biogenesis and provide material to rebuild destabilized network (Trogden et al. 2019). This is likely preparing cells for the next round of stimulation. Thus, existing data provide at least initial understanding of the mechanisms whereby β-cell MT architecture allows for fine-tuning of GSIS levels.

Revised version: During GSIS, the above-mentioned glucose-dependent destabilization of the sub-membrane MT array must reduce IG withdrawal and downplay the negative MT-dependent regulation of secretion(Bracey et al. 2020). At the same time, new MTs polymerizing off the Golgi facilitate IG biogenesis and provide material to rebuild destabilized network (Trogden et al. 2019). This is likely preparing cells for the next round of stimulation. Thus, existing data provide at least initial understanding of the mechanisms whereby β-cell MT architecture allows for fine-tuning of GSIS levels.

However, it is yet unclear how the complex β-cell MT network forms. Being nucleated at MTOCs in the cell interior, it is puzzling that the resulting β-cell MTs are not organized in conventional radial arrays and rather form a prominent peripheral array (Bracey et al. 2020). The goal of this study is to uncover the mechanisms underlying development of β-cell-specific MT network.

One of the established ways to modify the MT network without changing the location of MTOCs is to relocate already polymerized MTs by active motor-dependent transport. This phenomenon is called “MT sliding” (Straube et al. 2006). Several MT-dependent molecular motors have been implicated in driving MT sliding (Lu and Gelfand 2017). In some cases, a motor facilitates MT sliding by walking along a MT while its cargo-binding domain is stationary being attached to a relatively large structure, e.g. plasma membrane. This causes sliding of a MT which served as a track for the stationary motor. This mechanism has been described for dynein-dependent MT sliding (He et al. 2005, Grabham et al. 2007). MTs can also be efficiently slid by motors which have two functional motor assemblies, such as a tetrameric kinesin-5/Eg5 (Acar et al. 2013, Vukusic et al. 2021), or which carry a MT as a cargo while walking along another MT. For the latter mechanism, a motor needs a non-motor domain with a capacity to bind either a MT itself, or a MT-associated protein as an adapter (Kurasawa et al. 2004, Cao et al. 2020, Vukusic et al. 2021).

Out of these MT sliding factors, kinesin-1 is known to be critical for organizing unusual MT architecture in specialized cells. In oocytes, kinesin-1-dependent MT sliding empowers cytoplasmic streaming (Barlan et al. 2013). In differentiating neurons, kinesin-1 moves organelles and MTs into emerging neurites, which is a defining step in developing branched MT networks and long-distance neuronal transport (Jolly et al. 2010, Lu et al. 2013). With these data in mind, kinesin-1 presents itself as the most attractive candidate for organizing MTs in β cells. This motor is highly expressed in β cells and is well known to act as a major driving force in IG transport and GSIS (Meng et al. 1997, Varadi et al. 2002, Varadi et al. 2003, Cui et al. 2011).

Earlier version: Here, we show that KIF5B, being the major variant of kinesin 1 in β cells. actively slides MTs in this cell type. We show that this phenomenon defines MT network morphology and supplies MTs for the submembrane array. Moreover, we find that MT sliding in β cells is a glucose-dependent process and thus likely participates in metabolically driven cell reorganization during each secretion cycle.

Revised version: Here, we show that KIF5B, being the major variant of kinesin-1 in β cells, actively slides MTs in this cell type. We show that this phenomenon defines MT network morphology and supplies MTs for the sub-membrane array. Moreover, we find that MT sliding in β cells is a glucose-dependent process and thus likely participates in metabolically driven cell reorganization during each secretion cycle.

Results

Earlier version: Identification of KIF5B as MT-sliding motor in β cells

Revised version: Identification of KIF5B as the MT sliding motor in β cells

To address the factors that shape the configuration of MT networks in β cells, we tested for a potential involvement of motors-dependent MT sliding. Not surprisingly, analysis of existing RNA-sequencing data in functional mouse islet β cells highlighted kinesin-1 KIF5B as the highest expressing β-cell motor protein (Fig. 1A) (Sanavia et al. 2021). Since this kinesin has been reported to have MT sliding activity in many types of interphase cells, we tested its potential ability to slide MTs in β cells.

Earlier version: MTs in pancreatic beta cells undergo extensive sliding driven by kinesin KIF5B.

Revised version: MTs in pancreatic β cells undergo extensive sliding driven by kinesin KIF5B.

Earlier version: (A) A subset of RNA-sequencing data from primary mouse beta cells showing highly expressed kinesins as indicated by mRNA counts. KIF5B (most-right bar, red data points) is the most abundant kinesin motor in this cell type. N=3. Note that this is a subset of the RNA sequencing sets published in (Sanavia et al. 2021). (B) Efficient depletion of KIF5B in MIN6 cells using two alternative shRNA sequences, as compared to a scrambled shRNA control. Based on immunofluorescent staining of KIF5B as in Supplemental Figure 1 (A-C). Fold decrease of fluorescence signal per cell normalized to cells w/o shRNA expression in the same field of view. N= 25-32 cells from 4 repeats. (C) Quantification of MT sliding FRAP assay in cells treated with scrambled control or one of the two KIF5B-specific shRNAs (see representative data in D-F). MT displacement is shown as area of MTs displaced into the bleached area after 5 minutes of recovery. One-way ANOVA test was performed for statistical significance (p-value <0.0001). N=9-20 cells per set. (D-F) Frames from representative FRAP live-cell imaging sequences. mEmerald-tubulin-expressing MIN6 cells. Inverted grayscale images of maximum intensity projections of spinning disk confocal microscopy stacks over a 1 µm-thick ventral cell layer. Scale bars 5µm. (D1-F1) Overview of the whole cell prior to photobleaching. (D2-F3) Enlarged areas from (D1-F1) immediately after photobleaching (D2-F2) and 5 minutes (300 seconds) after photobleaching (D3-F3). Light-blue dotted lines indicate the edges of the photobleached areas. Red arrows indicate MTs displaced into the bleached area. Scale bars, 5 µm (D-F, Figure 1-Video 1 “MT Sliding FRAP”). (G-I) MIN6 cells featuring fiducial marks at MTs due to co-expression of SunTag-KIF5B-560Rigor construct and Halo-SunTag ligand. Representative examples for scrambled control shRNA-treated cell (G), KIF5B shRNA #1-treated cell (H) and KIF5B shRNA #2-treated cell (I) are shown. Single slices by spinning disk confocal microscopy. Halo-tag signal is shown as inverted gray-scale image. Top panels show cell overviews (scale bars 5µm). Below, boxed insets (scale bars 2 µm) are enlarged to show dynamics of fiducial marks (color arrows) at 1-second intervals (1-5 seconds). 0- to 5-second tracks of fiducial mark movement are shown in the bottom panel, each track color-coded corresponding to the arrows in the image sequences. (J) Summarized quantification of stationary fraction of fiducial marks along MT lattice (5-second displacements below 0.15µm). Scrambled shRNA control N=1,421 tracks across 6 cells, shRNA#1 N=852 tracks across 5 cells, shRNA#2 N=2,182 tracks across 7 cells. P One-way ANOVA, p<0.001 (K) Summarized quantification of motile fraction of fiducial marks along the MT lattice (5-second displacements above 0.3µm).. Scrambled shRNA control N=2,066 tracks across 6 cells, shRNA#1 N=390 tracks across 5 cells, shRNA#2 N=412 tracks across 7 cells. P One-way ANOVA, p<0.001(G-I Figure 1-Video 2 “MT Sliding SunTag”).

Revised version: (A) A subset of RNA-sequencing data from primary mouse β cells showing highly expressed kinesins as indicated by mRNA counts. KIF5B (most-right bar, red data points) is the most abundant kinesin motor in this cell type. N=3. Note that this is a subset of the RNA sequencing sets published in (Sanavia et al. 2021). (B) Efficient depletion of KIF5B in MIN6 cells using two alternative shRNA sequences, as compared to a scrambled shRNA control. Based on immunofluorescent staining of KIF5B as in Supplemental Figure 1 (A-C). Fold decrease of fluorescence signal per cell normalized to cells w/o shRNA expression in the same field of view. N= 25-32 cells from 4 repeats. (C) Quantification of MT sliding FRAP assay in cells treated with scrambled control or one of the two KIF5B-specific shRNAs (see representative data in D-F). MT displacement is shown as area of MTs displaced into the bleached area after 5 minutes of recovery. One-way ANOVA test was performed for statistical significance (p-value &lt;0.0001). N=9-20 cells per set. (D-F) Frames from representative FRAP live-cell imaging sequences. mEmerald-tubulin-expressing MIN6 cells. Inverted grayscale images of maximum intensity projections of spinning disk confocal microscopy stacks over a 1 µm-thick ventral cell layer. Scale bars 5µm. (D1-F1) Overview of the whole cell prior to photobleaching. (D2-F3) Enlarged areas from (D1-F1) immediately after photobleaching (D2-F2) and 5 minutes (300 seconds) after photobleaching (D3-F3). Light-blue dotted lines indicate the edges of the photobleached areas. Red arrows indicate MTs displaced into the bleached area. Scale bars, 5 µm (D-F, Figure 1-Video 1 “MT Sliding FRAP”). (G-I) MIN6 cells featuring fiducial marks at MTs due to co-expression of SunTag-KIF5B-560Rigor construct and Halo-SunTag ligand. Representative examples for scrambled control shRNA-treated cell (G), KIF5B shRNA #1-treated cell (H) and KIF5B shRNA #2-treated cell (I) are shown. Single slices by spinning disk confocal microscopy. Halo-tag signal is shown as inverted gray-scale image. Top panels show cell overviews (scale bars 5µm). Below, boxed insets (scale bars 2 µm) are enlarged to show dynamics of fiducial marks (color arrows) at 1-second intervals (1-5 seconds). 0-to 5-second tracks of fiducial mark movement are shown in the bottom panel, each track color-coded corresponding to the arrows in the image sequences. (J) Summarized quantification of stationary fraction of fiducial marks along MT lattice (5-second displacements below 0.15µm). Scrambled shRNA control N=1,421 tracks across 6 cells, shRNA#1 N=852 tracks across 5 cells, shRNA#2 N=2,182 tracks across 7 cells. P One-way ANOVA, p&lt;0.001 (K) Summarized quantification of motile fraction of fiducial marks along the MT lattice (5-second displacements above 0.3µm).. Scrambled shRNA control N=2,066 tracks across 6 cells, shRNA#1 N=390 tracks across 5 cells, shRNA#2 N=412 tracks across 7 cells. P One-way ANOVA, p&lt;0.001(G-I Figure 1-Video 2 “MT Sliding SunTag”).

Removed in the revised version: (A-C) shRNA-based depletion of KIF5B in MIN6 cells. Cells expressing shRNA are detected by GFP expression (cyan in A1-C1, red arrows in A2-C2). Immunostained KIF5B signal (magenta in A1-C1 and grayscale in A2-C2) is significantly reduced by KIF5B shRNAs (B, C) and compared to scrambled control (A). See quantification in main Figure 1B. (D) Histogram of displacements over 5-second intervals for all fiducial marks in scrambled shRNA control vs KIF5B shRNA #1 and KIF5B shRNA #2.

Efficient depletion of KIF5B was achieved by utilizing two independent lentiviral-based shRNA against mouse KIF5B in mouse insulinoma cell line MIN6 (Fig. 1B, Fig. 1-Suppl. Fig. 1A). To visualize MT sliding, shRNA-treated MIN6 cells expressing mEmerald-tubulin were imaged by live-cell spinning disk confocal microscopy. We photobleached MTs in two large cell regions leaving a thin unbleached band (“fluorescent belt”) and analyzed relocation of MTs from the “fluorescent belt” into the bleached areas over time. To minimize the effects of plausible MT polymerization and to reduce photobleaching, MTs were imaged for short time periods (5 mins). Strikingly, in control cells (treated with scrambled control shRNA) MTs were efficiently translocated from the “fluorescent belt” into the photobleached area, indicating that MT sliding events are prominent in this cell type (Fig. 1C,D). In contrast, MIN6 cells expressing either KIF5B shRNA variants displayed a significant loss of MT sliding ability (Fig. 1C,E,F, Fig.1-Video 1 “MT Sliding FRAP”), indicating that the loss of KIFB leads to the loss of MT sliding.

Earlier version: While the assay described above provides an easy visualization of MT sliding, it allows for visualization of only a subset of the MT network. To further corroborate the above findings, we used a less photodamaging system to visualize MT sliding that does not involve photobleaching and allows for evaluation of displacements within the whole MT network. To this end we applied a microtubule probe of fiducial marks, K560RigorE236A-SunTag (Tanenbaum et al. 2014) in MIN6 cells (Fig 1G-I). This probe contains the human kinesin-1 motor domain (residues 1–560) with a rigor mutation in the motor domain (K560RigorE236A) and fused to 24 copies of a GCN4 peptide. The rigor mutation in the motor domain causes it to bind irreversibly to microtubules (Rice et al. 1999). When co-expressed with a pHalo-tagged anti-GCN4 single-chain antibody (ScFv-GCN4-HaloTag-GB1-NLS), K560RigorE236A can recruit up to 24 of the Halo ligands to a single position on a microtubule. The pHalo-tagged anti GCN4 construct also contains a nuclear localization signal (NLS) which lends itself to reduce background of the unbound dye. This enables visualization of MT sliding events via single molecule tracking of the fiducial marks along the MT lattice, allowing us to analyze MT sliding behavior within the whole network with high temporal and spatial resolution (Fig. 1G-K, Fig. 1- Supplemental Fig. 1B, Fig.1-Video 2 “MT Sliding Suntag”).

Revised version: While the assay described above provides an easy visualization of MT sliding, it allows for visualization of only a subset of the MT network. To further corroborate the above findings, we used a less photodamaging system to visualize MT sliding that does not involve photobleaching and allows for evaluation of displacements within the whole MT network. To this end we applied a MT probe of fiducial marks, K560RigorE236A-SunTag (Tanenbaum et al. 2014) in MIN6 cells (Fig 1G-I). This probe contains the human kinesin-1 motor domain (residues 1–560) with a rigor mutation in the motor domain (K560RigorE236A) and fused to 24 copies of a GCN4 peptide. The rigor mutation in the motor domain causes it to bind irreversibly to MTs (Rice et al. 1999). When co-expressed with a pHalo-tagged anti-GCN4 single-chain antibody (ScFv-GCN4-HaloTag-GB1-NLS), K560RigorE236A can recruit up to 24 of the Halo ligands to a single position on a MT. The pHalo-tagged anti-GCN4 construct also contains a nuclear localization signal (NLS), which lends itself to reduce the background of the unbound dye. This enables visualization of MT sliding events via single molecule tracking of the fiducial marks along the MT lattice, allowing us to analyze MT sliding behavior within the whole network with high temporal and spatial resolution (Fig. 1G-K, Fig. 1-Supplemental Fig. 1B, Fig.1-Video 2 “MT Sliding Suntag”).

Earlier version: Our data indicate that in cells treated with scrambled control shRNA, a subset of K560RigorE236A-SunTag fiducial marks underwent rapid directional movements, interpreted as MT sliding events (Fig 1G, K). In contrast, in cells expressing KIF5B-specific shRNAs the fraction of rapidly moving fiducial marks was significantly reduced, while the fraction of stationary fiducial marks increased (Fig. 1 H, I, J), indicating the suppression of MT sliding.

Revised version: Our data indicate that in cells treated with scrambled control shRNA, a subset of K560RigorE236A-SunTag fiducial marks underwent rapid directional movements, interpreted as MT sliding events (Fig 1G, K). In contrast, in cells expressing KIF5B-specific shRNAs the fraction of rapidly moving fiducial marks was significantly reduced, while the fraction of stationary fiducial marks increased (Fig. 1 H, I, J), indicating the suppression of MT sliding. Collectively, these results indicate that KIF5B is necessary for MT sliding in MIN6 cells.

Removed in the revised version: Collectively, these results indicate that KIF5B is necessary for MT sliding in MIN6 cells.

KIF5B is required for β-cell MT organization

Because MT sliding mediated by KIF5B is a prominent phenomenon in β cells, we sought to test whether it has functional consequences for MT networks in these cells. Tubulin immunostaining revealed striking differences in MT organization between MIN6 cells treated with scrambled control shRNA versus KIF5B-specific shRNAs. While control cells had convoluted non-radial MTs with a prominent sub-membrane array, typical for β cells (Fig. 2A), KIF5B-depleted cells featured extra-dense MTs in the cell center and sparse receding MTs at the periphery (Fig. 2B, C). Significant reduction of tubulin staining intensity at the cell periphery (Fig. 2D) confirms the robustness of this phenotype in MIN6 cells.

Microtubule abundance and alignment at the cell periphery depend on KIF5B.

Earlier version: (A-C) MT organization in MIN6 cells expressing scrambled control shRNA (A), KIF5B-targeting shRNA #1 (B), or KIF5B-targeting shRNA #2 (C). Top, immunofluorescence staining for tubulin (grayscale, inverted). Bottom, immunofluorescence staining for KIF5B (cyan). Maximum intensity projection of 1µm at the ventral side of the cell. N=12. Scale bars: 5µm. (D) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-C). Mean values, black bars. One-way ANOVA, p<0.0001. N=7-9 cells. (E) Histograms of MT directionality within 1µm of cell boundary using perfected thresholds (see Figure 2 - Supplemental figure 3-4 for the analysis workflow and variants) in cells treated with scrambled control versus KIF5B-targeting shRNA. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed single confocal slices of shRNA-treated cell population immunostained for tubulin, as represented in (F-H). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to KIF5B depletions. Pixel numbers in the analysis: SCR N=106,780 pixels across 9 cells, shRNA #1 N=71,243 across 7 cells, shRNA #2 N= 60,087 across 7 cells. (F-H) Representative examples of MT directionality analysis in single confocal slices of shRNA-treated cells immunostained for tubulin, as quantified in (E). Single laser scanning confocal microscopy slices. (F) Scrambled control shRNA-treated cell. (G) KIF5B shRNA#1-treated cell. (H) KIF5B shRNA#1-treated cell. Overviews of cellular MT networks are shown as threshold to detect individual peripheral MTs (see Figure 2- Supplemental Figure 3, panel A5). (F1-H2) Directionality analysis outputs of regions from yellow boxes in (F-H) are shown color-coded for the angles between MTs and the nearest cell border. (I) Color code for (F1-H2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red.

Revised version: (A-C) MT organization in MIN6 cells expressing scrambled control shRNA (A), KIF5B-targeting shRNA #1 (B), or KIF5B-targeting shRNA #2 (C). Top, immunofluorescence staining for tubulin (grayscale, inverted). Bottom, immunofluorescence staining for KIF5B (cyan). Maximum intensity projection of 1µm at the ventral side of the cell. N=12. Scale bars: 5µm. (D) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-C). Mean values, black bars. One-way ANOVA, p&lt;0.0001. N=7-9 cells. (E) Histograms of MT directionality within 1µm of cell boundary using perfected thresholds (see Figure 2 - Supplemental figure 3-4 for the analysis workflow and variants) in cells treated with scrambled control versus KIF5B-targeting shRNA. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed single confocal slices of shRNA-treated cell population immunostained for tubulin, as represented in (F-H). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to KIF5B depletions. Pixel numbers in the analysis: SCR N=106,780 pixels across 9 cells, shRNA #1 N=71,243 across 7 cells, shRNA #2 N= 60,087 across 7 cells. (F-H) Representative examples of MT directionality analysis in single confocal slices of shRNA-treated cells immunostained for tubulin, as quantified in (E). Single laser scanning confocal microscopy slices. (F) Scrambled control shRNA-treated cell. (G) KIF5B shRNA#1-treated cell. (H) KIF5B shRNA#1-treated cell. Overviews of cellular MT networks are shown as threshold to detect individual peripheral MTs (see Figure 2 - Supplemental Figure 3, panel A5). (F1-H2) Directionality analysis outputs of regions from yellow boxes in (F-H) are shown color-coded for the angles between MTs and the nearest cell border. (I) Color code for (F1-H2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red.

Removed in the revised version: Microtubule abundance and alignment at the cell periphery depend on KIF5B in primary β cells in mouse islets.

Removed in the revised version: (A-B) MT organization in β cells within intact islets isolated from a wt (A) or KIF5B KO mice. Immunofluorescence staining for tubulin (grayscale, inverted). Maximum intensity projections of the laser scanning confocal microscopy stacks through the cell. N=10-11. Scale bars: 5µm. (C) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-B). Mean values, black bars. One-way ANOVA, p<0.0001. N=10-11 cells. (D-E) Representative examples of MT directionality analysis in single confocal slices of β cells from mouse islets immunostained for tubulin, as quantified in (F). Single laser scanning confocal microscopy slices. (D) Cells from wt mouse islets. (E) Cells from KIF5B KO mouse islets. Overviews of cellular MT networks are shown as threshold to detect individual peripheral MTs (see Figure 2- Supplemental Figure 3, panel A5). (D1-E2) Directionality analysis outputs of regions from yellow boxes in (D-E) are shown color-coded for the angles between MTs and the nearest cell border. (F) Color code for (D1-E2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red. (G) Histograms of MT directionality within 1µm of cell boundary using perfected thresholds (see Figure 2 - Supplemental figure 3-4 for the analysis workflow and variants) in β cells from wt versus KIF5B KO mice. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed single confocal slices of shRNA-treated cell population immunostained for tubulin, as represented in (D-E). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to KIF5B depletions. Pixel numbers in the analysis: wt N=180,709 pixels across 10 cells, KIF5B KO N=103,342 across 12 cells.

Removed in the revised version: Overexpression of truncated KIF5B motor lacking cargo binding does not affect MT density at β cell periphery.

Removed in the revised version: (A-D) MT organization in MIN6 cells expressing scrambled control shRNA (A-B) or KIF5B-targeting shRNA #2 (C-D). BFP-tagged kinesin-1 motor construct is ectopically expressed in (B, D). Top, immunofluorescence staining for tubulin (grayscale, inverted). Bottom, BFP immunofluorescence for motor expression (cyan). Maximum intensity projections of the laser scanning confocal microscopy stacks through the whole cell. N=8. Scale bars: 5µm.

Removed in the revised version: Workflow of MT directionality analysis.

Removed in the revised version: (A) Representation of the analysis workflow using a control DMSO-treated cell immunostained for tubulin as an example. (A1) An image of the original inverted grayscale confocal slice. (A2) A deconvolved image. (A3) Mask of the cell boundary. (A4) An image within the mask after application of standard % threshold. (A5) An image within the mask after application of a threshold optimizing detection of peripheral MTs for a particular cell. (A3-A5) are derivatives of (A2). (A6) A schematic illustrating map of distances from the nearest cell (mask) border per pixel (not to scale). (A7) A schematic illustrating map of angles per pixel (not to scale). A6 and A7 can be produced from A4 or A5. (A8) Color-coded output map of MT directionalities. A8 is a derivative of A7.

Removed in the revised version: Illustration of thresholding variations and their influence on the output analysis.

Removed in the revised version: (A) KIF5B-depleted cell as an example of contrast/threshold optimization for different MT densities and corresponding outputs. Upper row, contrast is adjusted for distinction of single MTs at the cell periphery. Lower row, contrast is adjusted for distinction of single MTs at the cell center. (A1) A deconvolved inverted grayscale confocal slice image contrasted to highlight peripheral MTs. (A2) The same deconvolved inverted grayscale confocal slice image contrasted to highlight central MTs. (A3) The image from A1 after application of a threshold optimized to highlight peripheral MTs. (A4) The image from A1 after application of a standard % threshold. (A5) Analysis outcome from (A3). (A6) Analysis outcome from (A4). (B) Histograms of MT directionality within 1um of cell boundary using standard thresholds (see Supplemental figure 3/1 for the analysis workflow) in cells treated with scrambled control versus KIF5B-targeting shRNA. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed shRNA-treated cell population, as represented in (F-H). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to KIF5B depletions. Pixel numbers in the analysis: SCR N=61,437 pixels across 9 cells, shRNA#1 N=15,215 pixels across 7 cells, shRNA#2 N= 21,125 pixels across 7 cells. (C) Histograms of MT directionality presented in this Figure panel B with depicted outliers. (D) Histograms of MT directionality presented in Figure 2 panel E with depicted outliers. (E) Histograms of MT directionality presented in Figure 2 – Supplemental Figure 1 panel G with depicted outliers.

Earlier version: We proceeded to investigate MT organization within primary β cells of the KIF5B KO mouse line (Kif5Bfl/−:RIP2-Cre, Fig. 2- Supplemental Fig. 1). These mice exhibit a β cell-specific deletion of Kif5B, as previously outlined (Cui et al., 2011). Our examination revealed a consistent presence of a significant peripheral array in the C57BL/6J control mice, while the KO counterparts exhibited a partial loss of this peripheral bundle. Specifically, the measured tubulin intensity at the cell periphery was significantly reduced in the KO mice compared to their wild-type counterparts (Fig. 2- Supplemental Fig. 1A-C). Thus, our results indicate that loss of KIF5B leads to a strong defect in MT location to the cell periphery location in both a β cell culture model and primary β cells within intact islets, which could be a direct consequence of disrupted sliding of MTs from the cell center to the periphery.

Revised version: We proceeded to investigate MT organization within primary β cells of the KIF5B KO mouse line (Kif5Bfl/−:RIP2-Cre, Fig. 2-Supplemental Fig. 1). These mice exhibit a β-cell-specific deletion of Kif5B, as previously outlined (Cui et al., 2011). Our examination revealed a consistent presence of a significant peripheral array in the C57BL/6J control mice, while the KO counterparts exhibited a partial loss of this peripheral bundle. Specifically, the measured tubulin intensity at the cell periphery was significantly reduced in the KO mice compared to their wild-type counterparts (Fig. 2-Supplemental Fig. 1A-C). Distinct extra-dense MTs were not common in islet β cells, possibly due to the 3D shape of a beta cell in tissue and/or compensatory mechanisms in organisms. Thus, our results indicate that loss of KIF5B leads to a strong defect in MT location to the cell periphery in both a β-cell culture model and primary β cells within intact islets, which could be a direct consequence of disrupted sliding of MTs from the cell center to the periphery.

Earlier version: An alternative explanation for the loss of peripheral MTs would be the loss of the mechanical modification of MT lattice by the force provided by moving kinesin motors. In cells with radial MT organization, this mechanism was shown to promote MT rescues and growth toward the cell periphery (Andreu-Carbo et al. 2022). To test the contribution of kinesin motor movement at MTs to MT density at the cell periphery, we have attempted to rescue the KIF5B-depletion phenotype by re-expressing kinesin-1 motor (see Fig. 3A-1), which is capable of moving along MT but is unable to bind MTs as cargos. Neither in KIF5B-depleted nor in wild-type MIN6 cells, peripheral MTs were affected by the motor re-expression (Fig. 2 Supplemental Fig. 2), indicating that kinesin-dependent MT growth does not noticeably contribute to the formation of peripheral MT array in this cell type.

Revised version: An alternative explanation for the loss of peripheral MTs would be the loss of the mechanical modification of the MT lattice by the force provided by moving kinesin motors. In cells with radial MT organization, this mechanism was shown to promote MT rescues and growth toward the cell periphery (Andreu-Carbo et al. 2022). To test the contribution of kinesin motor movement at MTs to MT density at the cell periphery, we have attempted to rescue the KIF5B-depletion phenotype by re-expressing kinesin-1 motor (see Fig. 3A-1), which is capable of moving along MT but is unable to bind MTs as cargos. Neither in KIF5B-depleted nor in wild-type MIN6 cells, peripheral MTs were affected by the motor re-expression (Fig. 2, Supplemental Fig. 2), indicating that kinesin-dependent MT growth does not noticeably contribute to the formation of the peripheral MT array in this cell type.

MT sliding is facilitated through the ATP-independent MT-binding domain of kinesin-1.

Earlier version: (A) Schematic of kinesin-1 (KIF5) and the Dominant Negative (KIFDN) and heterodimerization strategy. Top schematic shows full length KIF5s, consisting of the motor domain, stalk coil-coil domain and the tail. Three constructs utilized here include (1) The KIF5C motor domain tagged with a blue fluorescent protein (BFP) and the FRB for heterodimerization; (2) KIFDNwt construct with KIF5B Tail domain tagged with the mCherry fluorescent protein and the FKBP for heterodimerization. (3) KIFDNmut construct is the same as (2) but features a set of point mutations (magenta) making the ATP-independent MT-binding domain unable to bind MT lattice. (B) Quantification of MT sliding in FRAP assay in cells subjected to DN construct expression and heterodimerization. MT displacement is shown as area of MTs displaced into the bleached area after 5 minutes of recovery. See representative data (C-G). N= 5-25 per condition. One-way ANOVA test was performed for statistical significance (p-value <0.0001; ns, non-significant). (C-G”) Frames from representative FRAP live-cell imaging sequences of MIN6 cells expressing mEmerald-tubulin. Inverted grayscale images of maximum intensity projections over 1 µm-thick stacks by spinning disk confocal microscopy. (C1-G1) The first frame after photobleaching. (C2-G2) A frame 5 minutes (300 seconds) after photobleaching. Light-blue dotted lines indicate the edges of the photobleached areas. Red arrows indicate MTs displaced into the bleached area. Scale bars, 5 µm. (C3-G3) Schematics of experimental manipulation: green represents MTs, blue represents endogenous KIF5B, magenta represents KIFDNwt, purple represents KIFDNmut, gray represents KIF5C motor, orange bracket represents heterodimerizing agent (rap, rapalog). Conditions: (C1-C3) Untreated control. Only endogenous KIF5B is present. (D1-D3) KIFDNwt overexpression. Endogenous KIF5B is unable to bind MTs. (E1-E3) KIFDNmut overexpression. It does not bind MTs and does not interfere with endogenous KIF5B. (C-E, Figure 3-Video 1 “KFDN FRAP”) (F1-F3) KIFDNwt and KIF5C motor overexpression plus rapalog treatment. Heterodimerization creates a large pool of motors capable of MT sliding. (G-G”) KIFDNmut and KIF5C motor overexpression plus rapalog treatment. Heterodimerization creates a large pool of the motor non-functional in MT sliding (F-G, Figure 3-Video 2 “KFDN + Motor FRAP”).

Revised version: (A) Schematic of kinesin-1 (KIF5) and the Dominant Negative (KIFDN) and heterodimerization strategy. Top schematic shows full length KIF5s, consisting of the motor domain, stalk coil-coil domain and the tail. Three constructs utilized here include (1) The KIF5C motor domain tagged with a blue fluorescent protein (BFP) and the FRB for heterodimerization; (2) KIFDNwt construct with KIF5B Tail domain tagged with the mCherry fluorescent protein and the FKBP for heterodimerization. (3) KIFDNmut construct is the same as (2) but features a set of point mutations (magenta) making the ATP-independent MT-binding domain unable to bind MT lattice. (B) Quantification of MT sliding in FRAP assay in cells subjected to DN construct expression and heterodimerization. MT displacement is shown as area of MTs displaced into the bleached area after 5 minutes of recovery. See representative data (C-G). N= 5-25 per condition. One-way ANOVA test was performed for statistical significance (p-value &lt;0.0001; ns, non-significant). (C-G”) Frames from representative FRAP live-cell imaging sequences of MIN6 cells expressing mEmerald-tubulin. Inverted grayscale images of maximum intensity projections over 1 µm-thick stacks by spinning disk confocal microscopy. (C1-G1) The first frame after photobleaching. (C2-G2) A frame 5 minutes (300 seconds) after photobleaching. Light-blue dotted lines indicate the edges of the photobleached areas. Red arrows indicate MTs displaced into the bleached area. Scale bars, 5 µm. (C3-G3) Schematics of experimental manipulation: green represents MTs, blue represents endogenous KIF5B, magenta represents KIFDNwt, purple represents KIFDNmut, gray represents KIF5C motor, orange bracket represents heterodimerizing agent (rap, rapalog). Conditions: (C1-C3) Untreated control. Only endogenous KIF5B is present. (D1-D3) KIFDNwt overexpression. Endogenous KIF5B is unable to bind MTs. (E1-E3) KIFDNmut overexpression. It does not bind MTs and does not interfere with endogenous KIF5B. (C-E, Figure 3-Video 1 “KFDN FRAP”) (F1-F3) KIFDNwt and KIF5C motor overexpression plus rapalog treatment. Heterodimerization creates a large pool of motors capable of MT sliding. (G-G”) KIFDNmut and KIF5C motor overexpression plus rapalog treatment. Heterodimerization creates a large pool of the motor non-functional in MT sliding (F-G, Figure 3-Video 2 “KFDN + Motor FRAP”).

KIF5B is required for β-cell sub-membrane MT array alignment

Earlier version: Given the known significance of the peripheral MT array, which normally consists of well-organized MTs parallel to the cell membrane (Bracey et al. 2020), we have further analyzed directionality of MTs remaining at the cell periphery after KIF5B depletion. Previously we published a custom image analysis algorithm (Bracey et al. 2020) allowing for detailed quantitative characterization of MTs directionality in relation to the nearest cell border (Fig. 2 Supplemental Fig. 3). Here, we applied the same computational analysis to MT imaging data in MIN6 cells with perturbed KIF5B level and/or function. After deconvolution for increased signal-to-noise ratio, single 2D slices of MT images were subjected to thresholding. Two thresholding options were analyzed for unbiased approach. An threshold specifically optimized (perfect) for the peripheral MT array in each cell was considered along with a standard threshold across the analyzed cell population (Fig. 2 Supplemental Fig. 4). Thereafter, the directionality of MTs in respect to the cell border was determined. Every pixel of the image was analyzed with inconclusive pixels disregarded. MT directionality was quantified as a function of the distance from the cell border and directionality of peripheral MTs within 1µm of the cell border quantified. Across the analysis, two thresholding conditions provided similar outputs. Our results indicate that in MIN6 cells treated with non-targeting control shRNA (Fig. 2 F), the distribution of MT angles in the cell periphery are vastly parallel and co-aligned with the cell boundary, as previously reported for primary islet β cells (Bracey et al. 2020). In contrast, the loss of KIF5B via shRNA depletion resulted in a significant loss of parallel MTs at the periphery as compared to control (Fig. 2 E-H).

Revised version: Given the known significance of the peripheral MT array, which normally consists of well-organized MTs parallel to the cell membrane (Bracey et al. 2020), we have further analyzed directionality of MTs remaining at the cell periphery after KIF5B depletion. Previously we published a custom image analysis algorithm (Bracey et al. 2020) allowing for detailed quantitative characterization of MTs directionality in relation to the nearest cell border (Fig. 2 Supplemental Fig. 3). Here, we applied the same computational analysis to MT imaging data in MIN6 cells with perturbed KIF5B level and/or function. After deconvolution for increased signal-to-noise ratio, single 2D slices of MT images were subjected to thresholding. Two thresholding options were analyzed for the unbiased approach. A threshold specifically optimized (perfect) for the peripheral MT array in each cell was considered along with a standard threshold across the analyzed cell population (Fig. 2 Supplemental Fig. 4). Thereafter, the directionality of MTs with respect to the cell border was determined. Every pixel of the image was analyzed with inconclusive pixels disregarded. MT directionality was quantified as a function of the distance from the cell border and directionality of peripheral MTs within 1µm of the cell border quantified. Across the analysis, two thresholding conditions provided similar outputs. Our results indicate that in MIN6 cells treated with non-targeting control shRNA (Fig. 2 F), the distribution of MT angles in the cell periphery are vastly parallel and co-aligned with the cell boundary, as previously reported for primary islet β cells (Bracey et al. 2020). In contrast, the loss of KIF5B via shRNA depletion resulted in a significant loss of parallel MTs at the periphery as compared to control (Fig. 2 E-H).

Furthermore, the MT directionality analysis in β cells within islets from the KIF5B KO mice revealed a significant decrease in MT alignment to the cell periphery as compared to the wt mice, closely mirroring our findings from the KIF5B knockdown in MIN6 cells (Fig 2 Supplemental Fig. 2 D-G). Collectively, our data suggest that KIF5B-mediated MT sliding serves to facilitate MT transport to the cell periphery and align them at this cellular location in both a β-cell culture model and primary β cells within intact islets.

Earlier version: Combined, our data demonstrate a dramatic effect of KIF5B perturbation on both the distribution of MTs to the cell periphery and their orientation along the cell boundary. These data suggest that KIF5B-driven MT sliding is a decisive mechanism of the sub-membrane MT array generation, likely via redistribution of centrally nucleated MTs and subsequent aligning them at the cell edge. Thus, MT sliding is likely a critical component in functional MT organization in β cells.

Revised version: Combined, our data demonstrate a distinct effect of KIF5B perturbation on both the distribution of MTs to the cell periphery and their orientation along the cell boundary. These data suggest that KIF5B-driven MT sliding is a decisive mechanism of the sub-membrane MT array generation, likely via redistribution of centrally nucleated MTs and subsequent aligning them at the cell edge. Thus, MT sliding is likely a critical component in functional MT organization in β cells.

β-cell kinesin-1 drives MT sliding through the C-terminal MT-binding domain

While membrane cargo transport by KIF5s requires association of the heavy chain with the kinesin light chains (KLCs) and/or other adaptors, transportation of MTs as cargos occurs due to direct binding of KIF5 to MTs through the ATP-independent MT binding domain in heavy chain tail (C-terminus) (Jolly et al. 2010, Seeger and Rice 2010).

Earlier version: To specifically establish the role of MT sliding by KIF5B in β cells, we sought to evaluate the effects of suppressing the binding of KIF5B tail to MTs. To this end we used a previously generated construct (Ravindran et al. 2017), which is a motor-less version of wild-type (WT) kinesin-1 motor KIF5B containing the cargo-binding and ATP independent MT binding domain and tagged with mCherry (mCh) at the amino terminus (Fig 3A). When overexpressed, this construct acts as a dominant-negative (DN) tool preventing association of the tail of endogenous KIF5B with MTs. This tool is referred to as KIFDNwt (KIF5B dominant negative wild-type) moving forward (Fig. 3A).

Revised version: To specifically establish the role of MT sliding by KIF5B in β cells, we sought to evaluate the effects of suppressing the binding of KIF5B tail to MTs. To this end we used a previously generated construct (Ravindran et al. 2017), which is a motor-less version of wild-type (WT) kinesin-1 motor KIF5B containing the cargo-binding and ATP-independent MT binding domain and tagged with mCherry (mCh) at the amino terminus (Fig 3A). When overexpressed, this construct acts as a dominant-negative (DN) tool preventing the association of the tail of endogenous KIF5B with MTs. This tool is referred to as KIFDNwt (KIF5B dominant negative wild-type) moving forward (Fig. 3A).

To confirm that KIF5B tail domain binds to MTs in MIN6 and acts a dominant negative we co-expressed KIFDNwt and mEmerald-tubulin. When subjected to the FRAP assay we detected a complete loss of MT sliding events as compared to a control (Fig. 3B-D). To prevent tail engagement of the MT lattice through the ATP-independent binding domain we opted to make point mutations in the tail domain to change the residues 892-DRKRYQ to 892-DAAAYA, thus generating KIFDNMUT (Fig. 3A). Photobleaching assay in cells co-expressing of the KIFDNMUT with mEmerald-tubulin indicated that the MT sliding activity was not blocked in the presence of the mutated construct (Fig. 3B, E, Fig.3-Video 1 KFDN FRAP), confirming that KIF5B tail domain binding to MTs is needed for MT sliding in β cells.

Earlier version: The dominant negative constructs are also tagged with the FK506-rapamycin-binding protein (FKBP), as indicated in Fig. 3A. This allows to heterodimerize them with a motor domain fused with the FKBP-rapamycin binding (FRB) domain using A/C Heterodimerizer (rapalog) and reconstitute a functional motor (Inobe and Nukina 2016). We restored kinesin-1 activity by connecting the motor-less KIF5B, KIFDNwt, to kinesin-1 motor domain as a way to rescue the effects of DN approach of KIF5B tail overexpression. To this end, we co-expressed MIN6 cells with the tail domain, mEmerald-tubulin, and the KIF5C motor domain fused to FRB domain (Fig. 3A). Once the tail and motor domain were dimerized with rapalog, we saw that the once blocked MT sliding events of the KIFDNwt tail alone were now reversed (Fig. 3B, F, Fig.3-Video 2 KFDN + Motor FRAP). In contrast, under conditions of heterodimerization of KIFDNMUT with the motor, MT sliding was greatly impaired (Fig. 3B, G, Fig.3-Video 2 KFDN + Motor FRAP), indicating that the motor with mutated ATP-independent binding domain cannot use MTs as cargos.

Revised version: The dominant negative constructs are also tagged with the FK506-rapamycin-binding protein (FKBP), as indicated in Fig. 3A. This allows to heterodimerize them with a motor domain fused with the FKBP-rapamycin binding (FRB) domain using A/C Heterodimerizer (rapalog) and reconstitute a functional motor (Inobe and Nukina 2016). We restored kinesin-1 activity by connecting the motor-less KIF5B, KIFDNwt, to kinesin-1 motor domain as a way to rescue the effects of DN approach of KIF5B tail overexpression. To this end, we co-expressed MIN6 cells with the tail domain, mEmerald-tubulin, and the KIF5C motor domain fused to FRB domain (Fig. 3A). Once the tail and motor domain were dimerized with rapalog, we saw that the once blocked MT sliding events of the KIFDNwt tail alone were now reversed (Fig. 3B, F, Fig.3-Video 2 KFDN + Motor FRAP). In contrast, under conditions of heterodimerization of KIFDNMUT with the motor, MT sliding was greatly impaired (Fig. 3B, G, Fig.3-Video 2 KFDN + Motor FRAP), indicating that the motor with the mutated ATP-independent binding domain cannot use MTs as cargos. Interestingly, the endogenous motor in this case was unable to efficiently transport MTs, suggesting that the endogenous motor pool engaged in MT sliding was significantly smaller than the overexpressed non-functional motor.

Removed in the revised version: Interestingly, endogenous motor in this case was unable to efficiently transport MTs, suggesting that the endogenous motor pool engaged in MT sliding was significantly smaller than the overexpressed non-functional motor.

Overall, the results of the DN approach confirm that MT sliding in β cells is driven by KIF5B through direct kinesin-1 tail binding to cargo MTs.

Effects of C-terminal MT-binding of kinesin-1 on β-cell MT organization

Keeping in mind that KIF5B has additional major functions in addition to MT sliding, we sought to test the consequence of MT sliding more directly by turning to overexpression of the DN constructs. Thus, we took advantage of our heterodimerization approach to analyze MT patterns in cells with active kinesin-1, which is able or unable to slide MTs (see Fig.3F vs Fig.3G). We analyzed MIN6 cells that express either the KIFDNwt or KIFDNmut tail domains alone (Fig. 4 Supplemental Fig. 1) or co-expressing and heterodimerized with the motor domain (Fig. 4B,C). Cells were fixed and immunostained for tubulin to identify the MT network. As expected, overexpression of the KIFDNWT tail construct alone acted as dominant negative toward MT distribution to the cell periphery, resulting in decreased peripheral tubulin intensity (Fig. 4 Supplemental Fig. 1A,C), while in cells expressing KIFDNMUT, MT patterns were comparable to control (Fig. 4 Supplemental Fig. 1B,C).

Earlier version: Effects of ATP-independent MT-binding domain of KIF5B on microtubule abundance and alignment at the β-cell cell periphery.

Revised version: Effects of ATP-independent MT-binding domain of KIF5B on MT abundance and alignment at the β-cell cell periphery.

Earlier version: (A-C) MT organization in MIN6 cells expressing (B) KIFDNwt and KIF5C motor heterodimerized via rapalog treatment, (C) KIFDNmut and KIF5C motor heterodimerized via rapalog treatment and compared to a control cell with no ectopic expressions (A). Top, immunofluorescence staining for tubulin (grayscale, inverted). Blue dotted line indicates the borders of a cell expressing constructs of interest. Bottom in A, f-actin (phalloidin, red) and DAPI (blue). Bottom in B and C, ectopically expressed mCherry-labeled KIFDN constructs (magenta) and BFP-labeled KIF5C motor (green). Laser scanning confocal microscopy maximum intensity projection of 1µm at the ventral side of the cell. Scale bars: 5um. (D) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-C). Mean values, black bars. One-way ANOVA, p<0.0001. N=7-15 cells. (E) Histograms of MT directionality within 1um of cell boundary (see Supplemental figure 2-1 for the analysis workflow) in control cells compared to cells expressing heterodimerized KIFDN variants. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed cell population, as represented in (F-H). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to the over-expressions. NT control N=138,810 pixels across 9 cells, KIFDNwt + motor N= 48,285 pixels across 9 cells, KIFDNmut + motor N= 40,832 pixels across 10 cells. (F-H) Representative examples of MT directionality analysis quantified in (E). (F) Control cell, no ectopic expressions. (G) Cell expressing KIFDNwt+ Motor. (H) Cell expressing KIFDNmut+ Motor. Overviews of cellular MT networks are shown as threshold to detect individual peripheral MTs (see Supplemental figure 2-1 panel A5). (F1-H2) Directionality analysis outputs of regions from yellow boxes in (F-H) are shown color-coded for the angles between MTs and the nearest cell border (see Supplemental figure 2-1 panel A8). (I) Color code for (F1-H2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red.

Revised version: (A-C) MT organization in MIN6 cells expressing (B) KIFDNwt and KIF5C motor heterodimerized via rapalog treatment, (C) KIFDNmut and KIF5C motor heterodimerized via rapalog treatment and compared to a control cell with no ectopic expressions (A). Top, immunofluorescence staining for tubulin (grayscale, inverted). Blue dotted line indicates the borders of a cell expressing constructs of interest. Bottom in A, f-actin (phalloidin, red) and DAPI (blue). Bottom in B and C, ectopically expressed mCherry-labeled KIFDN constructs (magenta) and BFP-labeled KIF5C motor (green). Laser scanning confocal microscopy maximum intensity projection of 1µm at the ventral side of the cell. Scale bars: 5um. (D) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-C). Mean values, black bars. One-way ANOVA, p&lt;0.0001. N=7-15 cells. (E) Histograms of MT directionality within 1um of cell boundary (see Supplemental figure 2-1 for the analysis workflow) in control cells compared to cells expressing heterodimerized KIFDN variants. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed cell population, as represented in (F-H). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to the over-expressions. NT control N=138,810 pixels across 9 cells, KIFDNwt + motor N= 48,285 pixels across 9 cells, KIFDNmut + motor N= 40,832 pixels across 10 cells. (F-H) Representative examples of MT directionality analysis quantified in (E). (F) Control cell, no ectopic expressions. (G) Cell expressing KIFDNwt+ Motor. (H) Cell expressing KIFDNmut+ Motor. Overviews of cellular MT networks are shown as threshold to detect individual peripheral MTs (see Supplemental figure 2-1 panel A5). (F1-H2) Directionality analysis outputs of regions from yellow boxes in (F-H) are shown color-coded for the angles between MTs and the nearest cell border (see Supplemental figure 2-1 panel A8). (I) Color code for (F1-H2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red.

Removed in the revised version: DN effect of KIF5B tail domain on MT abundance at the β- cell cell periphery requires ATP-independent MT-binding domain.

Removed in the revised version: (A-B) MT organization in MIN6 cells expressing (A) KIFDNwt, (B) KIFDNmut, Top, immunofluorescence staining for tubulin (grayscale, inverted). Blue dotted line indicates the borders of a cell expressing constructs of interest. Bottom, ectopically expressed mCherry-labeled KIFDN constructs (magenta). Laser scanning confocal microscopy maximum intensity projection of 1µm at the ventral side of the cell. Scale bars: 5um. (C) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-B) as compared to untreated controls (see Figure 4A). Mean values, black bars. One-way ANOVA, p<0.0001. N=7-11 cells.

Removed in the revised version: Influence of thresholding variations on the output analysis of MT directionality in Figure 4 F-H.

Removed in the revised version: (A) Histograms of MT directionality within 1um of cell boundary using standard thresholds (see Supplemental figure 2-1 for the analysis workflow) in cells expressing ectopic heterodimerized KIF5C motor. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed shRNA-treated cell population, as represented in (Figure 4 F-H). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to heterodimerized KIF5B over-expression. Pixel numbers in the analysis: NT control N=120,959 pixels across 9 cells, KIFDNwt + motor N= 26,110 pixels across 7 cells, KIFDNmut + motor N= 41,091 pixels across 10 cells. (B) Histograms of MT directionality presented in this Figure panel A with depicted outliers. (C) Histograms of MT directionality presented in main Figure 4 panel E with depicted outliers.

Interestingly, expression of heterodimerized kinesin motors led to impaired MT network configurations compared to NT control (Fig. 4). Specifically, blocking of MT sliding by overexpression KIFDNmut heterodimerized with the motor, led to the decrease in peripheral tubulin intensity (Fig. 4C,D) and impaired MT aligning along the cell border (Fig. 4E,G). These data indicate that KIF5B-driven relocation of centrally nucleated MTs to β-cell periphery requires kinesin tail domain binding to “cargo” MTs. Strikingly, overexpression of functional heterodimerized motor, which was capable of MT sliding and populating of the cell periphery with MTs as detected by tubulin intensity readings (KIFDNwt heterodimerized with the motor, Fig. 4B,D), also led to a deficient MT aligning at the periphery (Fig. 4E,H). This can be interpreted as a result of unregulated sliding in these experimental conditions, since excessive kinesin-1-dependent sliding can lead to MT bending (Straube et al. 2006). This suggests that proper organization of MTs within the sub-membrane array requires fine tuning of MT sliding activity.

Collectively, this indicates that regulated KIF5B activity is essential for redistributing MTs to the cell border and sustaining an aligned peripheral MT array.

Exaggerated MT sliding leads to defects in peripheral array alignment

Earlier version: Our data discussed above suggest that overexpression of functional kinesin-1 disrupts MT alignment at the cell periphery, inducing their bending and buckling (Fig. 4E, G). To test if this defect is a result of excessive MT sliding, we employed a small molecule, kinesore, which is known to dramatically promote MT sliding by kinesisn-1 (Randall et al. 2017). Kinesore targets kinesin cargo adaptor function, by impairing KLC from binding kinesin heavy chain. As a result, kinesin heavy chain will excessively engage MTs through the C-terminal, ATP independent MT binding domain, leading to exaggerated MT sliding and the loss of membrane cargo transport by kinesin-1 (Randall et al. 2017).To this end we pretreated MIN6 cells with 50µm kinesore and stained for MTs (Fig. 5).

Revised version: Our data discussed above suggest that overexpression of functional kinesin-1 disrupts MT alignment at the cell periphery, inducing their bending and buckling (Fig. 4E, G). To test if this defect is a result of excessive MT sliding, we employed a small molecule, kinesore, which is known to dramatically promote MT sliding by kinesin-1 (Randall et al. 2017). Kinesore targets kinesin cargo adaptor function, by impairing KLC from binding kinesin heavy chain. As a result, kinesin heavy chain will excessively engage MTs through the C-terminal, ATP independent MT binding domain, leading to exaggerated MT sliding and the loss of membrane cargo transport by kinesin-1 (Randall et al. 2017).To this end we pretreated MIN6 cells with 50µm kinesore and stained for MTs (Fig. 5).

Enhanced MT sliding results in loss of peripheral MT alignment at the border.

Earlier version: (A-B) MT organization in MIN6 cells pretreated with DMSO and Kinesore respectively. Immunofluorescence staining for tubulin (grayscale, inverted). Laser scanning confocal microscopy maximum intensity projection of 1µm at the ventral side of the cell. Scale bars: 5um. (C) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-B). Mean values, black bars. One-way ANOVA, p<0.0001. N=10 cells per condition. Histograms of MT directionality within 1um of cell boundary (see Supplemental figure 2-1 for the analysis workflow) in DMSO treated control cells compared to kinesore treated cells. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed cell population, as represented in (D-E). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to the over-expressions. DMSO control N=136,840 pixels across 10 cells, kinesore treated N= 87,361 pixels across 9 cells. (D-E) Representative examples of MT directionality analysis quantified in (F). Directionality analysis outputs of regions from yellow boxes in (D-E) are shown color-coded for the angles between MTs and the nearest cell border (see Supplemental figure 2-1 panel A8). (G) Color code for (D1-E2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red.

Revised version: (A-B) MT organization in MIN6 cells pretreated with DMSO and Kinesore respectively. Immunofluorescence staining for tubulin (grayscale, inverted). Laser scanning confocal microscopy maximum intensity projection of 1µm at the ventral side of the cell. Scale bars: 5um. (C) Quantification of mean tubulin intensity within the outer 2µm peripheral area of a cell, in data represented in (A-B). Mean values, black bars. One-way ANOVA, p&lt;0.0001. N=10 cells per condition. Histograms of MT directionality within 1um of cell boundary (see Supplemental figure 2-1 for the analysis workflow) in DMSO treated control cells compared to kinesore treated cells. Data are shown for the summarized detectable tubulin-positive pixels in the analyzed cell population, as represented in (D-E). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to the over-expressions. DMSO control N=136,840 pixels across 10 cells, kinesore treated N= 87,361 pixels across 9 cells. (D-E) Representative examples of MT directionality analysis quantified in (F). Directionality analysis outputs of regions from yellow boxes in (D-E) are shown color-coded for the angles between MTs and the nearest cell border (see Supplemental figure 2-1 panel A8). (G) Color code for (D1-E2): MTs parallel to the cell edge, blue; MTs perpendicular to the cell edge, red.

Removed in the revised version: Influence of thresholding variations on the output analysis of MT directionality in Figure 5 D-E.

Removed in the revised version: (A) Histograms of MT directionality within 1um of cell boundary using standard thresholds (see Supplemental figure 2-1 for the analysis workflow) in cells treated with DMSO or kinesore. Data are shown for the summarized detectable tubulin- positive pixels in the analyzed respective cell population, as represented in (Figure 5 D-E). Unpaired t-test were performed across each bin for all cells, and a K-S test was performed on the overall distribution. The share of MTs parallel to the edge (bin 0-10) is significantly higher in control as compared to kinesore treated. Pixel numbers in the analysis: DMSO control N= 94,841 pixels across 9 cells, kinesore N= 29,796 pixels across 9 cells. (B) Histograms of MT directionality presented in this Figure panel A with depicted outliers. (C) Histograms of MT directionality presented in main Figure 5 panel F with depicted outliers.

Removed in the revised version: Kinesore has no effect in MT network in KIF5B-depleted cells.

Removed in the revised version: (A-F) Representative examples of MT organization in MIN6 cells expressing scrambled control shRNA (A, D), KIF5B-targeting shRNA #1 (B, E), or KIF5B-targeting shRNA #2 (C, F). A DMSO-treated control (A-C) and kinesore-treated cells (D-F) are shown. Immunofluorescence staining for tubulin (grayscale, inverted). Maximum intensity projections of the laser scanning confocal microscopy stacks through the whole cell. N=8 cells/condition. Scale bars: 5µm.

We have observed an exaggerated MT looping resulting in a slight decrease of peripheral MT intensity (Fig. 5C). To validate whether the effect of kinesore on MT morphology was KIF5B-dependent, we applied kinesore treatment to cells treated with either scrambled control or KIF5B-specific shRNA (Fig. 5 Supplemental Fig. 1). MT organization in KIF5B-depleted cells was not affected by kinesore, indicating that the observed effect was likely due to kinesore-induced MT sliding exaggeration (Fig. 5 Supplemental Fig. 2). Further analysis of the peripheral bundle indicated that MT alignment was strongly impaired upon kinesore-driven MT remodeling as compared with vehicle (DMSO) treatment (Fig. 5D-G). The loss of coaligned MTs and loss of tubulin density at the periphery indicate that MT sliding must be gated to prevent over-corrected MT networks.

Earlier version: MT sliding in β cells is activated by glucose stimulations

Revised version: MT sliding in β cells is activated by glucose stimulation

Earlier version: It has previously been reported that kinesin-1 switches activity level in the presence of glucose stimuli (Donelan et al. 2002). We predicted that as KIF5B activity modulates the MT sliding events would also change depending on the glucose concentration. To test this, we pre-incubated MIN6 cells with media containing a low concentration of 2.8mM glucose (Fig. 6A). We applied the photobleaching assay at these conditions and detected little to no MT sliding events. When switching glucose to a high concentration of 20mM, MT sliding and remodeling events were significantly increased (Fig. 6B, Fig.6-Video 1 “FRAP Low and High Glucose”).

Revised version: It has previously been reported that kinesin-1 switches activity level in the presence of glucose stimuli (Donelan et al. 2002). We predicted that, as KIF5B activity modulates the MT sliding events, they would also change depending on the glucose concentration. To test this, we pre-incubated MIN6 cells with media containing a low concentration of 2.8mM glucose (Fig. 6A). We applied the photobleaching assay under these conditions and detected little to no MT sliding events. When switching glucose to a high concentration of 20mM, MT sliding and remodeling events were significantly increased (Fig. 6B, Fig.6-Video 1 “FRAP Low and High Glucose”). Quantification of the sliding events demonstrated that MIN6 displaced MTs via MT sliding significantly more efficiently upon glucose stimulation (Fig. 6C). We then turned to single-molecule tracking of MT lattice fiducial marks (K560RigorE236A-SunTag) to further investigate this observation. Consistent with the photobleaching assay, the fiducial marks were predominantly stationary in cells pre-incubated in 2.8mM glucose (Fig. 6D) but frequently underwent directed relocation events indicative of MT sliding in cells after being stimulated with 20mM glucose (Fig. 6E, Fig.6-Video 2 “SunTag Low and High Glucose).

MT sliding in β cells is stimulated by glucose.

Earlier version: (A-B) Frames from representative FRAP live-cell imaging sequences of MT sliding response to glucose stimulation. mEmerald-tubulin-expressing MIN6 cells. Inverted grayscale images of maximum intensity projections over 1 µm-thick stacks by spinning disk confocal microscopy. (A) A cell pretreated with 2.8mM glucose before the assay. (B) A cell pretreated with 2.8mM glucose and stimulated with 20 mM glucose before the assay. (A1-B1) The first frame after photobleaching. (A2-B2) A frame 5 minutes (300 seconds) after photobleaching. Light-blue dotted lines indicate the edges of the photobleached areas. Red arrows indicate MTs displaced into the bleached area. Scale bars, 5 µm. (C) Quantification of MT sliding FRAP assay in cells in 2.8mM versus 20mM glucose (see representative data in A-B). MT displacement is shown as area of MTs displaced into the bleached area after 5 minutes of recovery. One-way ANOVA test was performed for statistical significance (p-value <0.0001). N=16-24 cells per set (A-B Figure 6-Video 1 “ FRAP Low and High Glucose”). (D-E) MIN6 cells featuring fiducial marks at MTs due to co-expression of SunTag-KIF5B-560Rigor construct and Halo-SunTag ligand. Representative examples for cells in 2.8mM glucose (D) and a cell stimulated by 20mM glucose (E) are shown. Single-slice spinning disk confocal microscopy. Halo-tag signal is shown as inverted gray-scale image. Top panels show cell overviews (scale bars 5µm). Below, boxed insets are enlarged to show dynamics of fiducial marks (color arrows) at 1 second intervals (1-5 seconds). 0- to 5-second tracks of fiducial mark movement are shown in the bottom panel, each track color-coded corresponding to the arrows in the image sequences. N=6-11 cells (A-B Figure 6-Video 2 “ SunTag Low and High Glucose”). (F) Histogram of all 5-second displacement of fiducial marks in Low vs High glucose (G) Summarized quantification of stationary fiducial marks along MT lattice (5-second displacements below 0.15µm). Low glucose N=5,615 tracks across 11 cells, High glucose 5min N=2,259 tracks across 6 cells, High Glucose 20min N=3,059 tracks across 6 cells. P One-way ANOVA, p<0.001 (H) Summarized quantification of moving fiducial marks along the MT lattice (5-second displacements over 0.3µm). Low glucose N=2,595 tracks across 11 cells, High glucose 5min N=2,642 tracks across 6 cells, High Glucose 20min N=4,049 tracks across 6 cells. One-way ANOVA, p<0.001.

Revised version: (A-B) Frames from representative FRAP live-cell imaging sequences of MT sliding response to glucose stimulation. mEmerald-tubulin-expressing MIN6 cells. Inverted grayscale images of maximum intensity projections over 1 µm-thick stacks by spinning disk confocal microscopy. (A) A cell pretreated with 2.8mM glucose before the assay. (B) A cell pretreated with 2.8mM glucose and stimulated with 20 mM glucose before the assay. (A1-B1) The first frame after photobleaching. (A2-B2) A frame 5 minutes (300 seconds) after photobleaching. Light-blue dotted lines indicate the edges of the photobleached areas. Red arrows indicate MTs displaced into the bleached area. Scale bars, 5 µm. (C) Quantification of MT sliding FRAP assay in cells in 2.8mM versus 20mM glucose (see representative data in A-B). MT displacement is shown as area of MTs displaced into the bleached area after 5 minutes of recovery. One-way ANOVA test was performed for statistical significance (p-value &lt;0.0001). N=16-24 cells per set (A-B Figure 6-Video 1 “FRAP Low and High Glucose”). (D-E) MIN6 cells featuring fiducial marks at MTs due to co-expression of SunTag-KIF5B-560Rigor construct and Halo-SunTag ligand. Representative examples for cells in 2.8mM glucose (D) and a cell stimulated by 20mM glucose (E) are shown. Single-slice spinning disk confocal microscopy. Halo-tag signal is shown as inverted gray-scale image. Top panels show cell overviews (scale bars 5µm). Below, boxed insets are enlarged to show dynamics of fiducial marks (color arrows) at 1 second intervals (1-5 seconds). 0-to 5-second tracks of fiducial mark movement are shown in the bottom panel, each track color-coded corresponding to the arrows in the image sequences. N=6-11 cells (A-B Figure 6-Video 2 “SunTag Low and High Glucose”). (F) Histogram of all 5-second displacement of fiducial marks in Low vs High glucose (G) Summarized quantification of stationary fiducial marks along MT lattice (5-second displacements below 0.15µm). Low glucose N=5,615 tracks across 11 cells, High glucose 5min N=2,259 tracks across 6 cells, High Glucose 20min N=3,059 tracks across 6 cells. P One-way ANOVA, p&lt;0.001 (H) Summarized quantification of moving fiducial marks along the MT lattice (5-second displacements over 0.3µm). Low glucose N=2,595 tracks across 11 cells, High glucose 5min N=2,642 tracks across 6 cells, High Glucose 20min N=4,049 tracks across 6 cells. One-way ANOVA, p&lt;0.001.

Removed in the revised version: Quantification of the sliding events demonstrated that MIN6 displaced MTs via MT sliding significantly more efficient upon glucose stimulation (Fig. 6C). We then turned to single molecule tracking of MT lattice fiducial marks (K560RigorE236A-SunTag) to further investigate this observation. Consistent with the photobleaching assay, the fiducial marks were predominantly stationary in cells pre-incubated in 2.8mM glucose (Fig. 6D) but frequently underwent directed relocation events indicative of MT sliding in cells after stimulated with 20mM glucose (Fig. 6E, Fig.6-Video 2 “SunTag Low and High Glucose).

These data demonstrate that glucose-stimulated remodeling of the MT network involves regulated MT sliding. Given the importance of MT sliding for peripheral MT abundance (Fig. 7A) and alignment (Figs. 7B), this effect may be essential to remodel parts of the peripheral MT array during GSIS to allow for secretion (Fig. 7C) and/or to restore the initial array after glucose-dependent destabilization (Fig. 7D).

Schematic of the main results and predictions.

(A) Role of KIF5B-dependent MT sliding in MT density at the cell periphery. (B) Role of KIF5B-dependent MT sliding in the alignment of peripheral MTs. (C) Potential role of glucose-facilitated KIF5B-dependent MT sliding in forming temporal openings in the peripheral MT array by moving fragmented MTs aside or MT looping. (D) Potential role of glucose-facilitated KIF5B-dependent MT sliding in repair of the peripheral MT array, partially destabilized/fragmented downstream of glucose. Track MTs, green. Cargo MTs, blue. Kinesin-1, black. Direction of kinesin movement shown as solid arrow. Subsequent steps of the process shown as dashed double arrows.

Discussion

Earlier version: Since the first description of convoluted MT network in MIN6 cells by the Rutter group (Varadi et al. 2003), our views on regulation, function, and dynamics of pancreatic β cell MT network have been gradually evolving (Bracey et al. 2022). However, the field is still far from the understanding of the mechanisms underlying the network architecture. Here, we show that MT sliding is a prominent phenomenon in β cells, and that it is driven by kinesin KIF5B. This kinesin-1-dependent MT sliding is a critical mechanism needed to formation and a long-term maintenance of β cell MT networks. In addition, we show that MT sliding activity is facilitated by glucose stimulation, suggesting that this process is involved in regulation of GSIS and/or providing β cell fitness during the response to glucose. Overall, our study establishes MT sliding as an essential regulator of β cell architecture and function.

Revised version: Since the first description of the convoluted MT network in MIN6 cells by the Rutter group (Varadi et al. 2003), our views on the regulation, function, and dynamics of the pancreatic β-cell MT network have been gradually evolving (Bracey et al. 2022). However, the field is still far from understanding the mechanisms underlying the network architecture. Here, we show that MT sliding is a prominent phenomenon in β cells and that it is driven by kinesin KIF5 B. This kinesin-1-dependent MT sliding is a critical mechanism needed for the formation and long-term maintenance of β-cell MT networks. In addition, we show that MT sliding activity is facilitated by glucose stimulation, suggesting that this process is involved in the regulation of GSIS and/or providing β-cell fitness during the response to glucose. Overall, our study establishes MT sliding as an essential regulator of β-cell architecture and function.

Earlier version: As we reviewed in the introduction, β cell MT network consists of interlocked meshwork and peripheral MT arrays co-aligned with the cell border. We find that blocking kinesin has two major effects of on MT organization: one, receding of MTs at the cell periphery and increased MT density in the cell center and two, lack of alignment of remaining peripheral MTs with the cell border. We interpret that the first phenotype arises from the lack of MT sliding from the cell center to periphery and the second phenotype arises from the lack of MT sliding along pre-existing peripheral MTs (Fig. 7).

Revised version: As we reviewed in the introduction, the β-cell MT network consists of an interlocked meshwork and peripheral MT arrays co-aligned with the cell border. We find that blocking kinesin has two major effects on MT organization: one, receding of MTs at the cell periphery and increased MT density in the cell center, and two, lack of alignment of remaining peripheral MTs with the cell border. We interpret that the first phenotype arises from the lack of MT sliding from the cell center to the periphery, and the second phenotype arises from the lack of MT sliding along pre-existing peripheral MTs (Fig. 7).

Earlier version: In cultured mesenchymal cells, kinesin motor was reported to populate cell periphery with MTs via promoting rescues and elongation of radially-arranged dynamic MTs (Andreu-Carbo et al. 2022). Our data indicate that this mechanism does not noticeably contribute to kinesin-dependent MT organization in β cells, leaving MT sliding the only known mechanism underlying the phenotype observed here. This difference in kinesin action could be possibly due to basic differences on MT network properties in these cell types: in a mostly non-radial highly-stabilized MT network in β cells (Zhu et al. 2015), modulation of MT plus end rescue efficiency in not likely to be a significant factor.

Revised version: In cultured mesenchymal cells, the kinesin motor was reported to populate the cell periphery with MTs via promoting rescues and elongation of radially-arranged dynamic MTs (Andreu-Carbo et al. 2022). Our data indicate that this mechanism does not noticeably contribute to kinesin-dependent MT organization in β cells, leaving MT sliding the only known mechanism underlying the phenotype observed here. This difference in kinesin action could be due to basic differences in MT network properties in these cell types: in a mostly non-radial, highly-stabilized MT network in β cells (Zhu et al. 2015), modulation of MT plus end rescue efficiency is not likely to be a significant factor.

Removed in the revised version: Interestingly, blocking kinesin results in a striking accumulation of MT in the cell center where they are normally nucleated at MTOCs, which include the centrosome and the Golgi, in differentiated β cells the latter being the main MTOC. Thus, sliding MTs originate from the MTOC area. At the same time, FIB-SEM analysis did not detect many MTs associated with MTOCs in physiologically normal β cells (Muller et al. 2021). This implicates that MTs are normally rapidly dissociated from MTOCs so that they become available for transport by sliding. It is worth mentioning that for long-distance transport by sliding, cargo MTs must be short, otherwise MT buckling and not long-distance transport will occur (Straube et al. 2006).

Earlier version: Interestingly, shorter MTs have been observed in high glucose conditions (Muller et al. 2021) when MT are nucleated more actively (Trogden et al. 2019) and transported more frequently (this paper). Possibly, nucleated MTs are detached from MTOCs before they achieve a length that would prevent their transport. An intriguing possibility has been proposed, indicating that in high glucose conditions, microtubules (MTs) might undergo severing by katanin (Muller et al. 2021). This process could generate MT fragments, potentially facilitating their role as cargos with increased ease of transport. It is also possible that sliding MT subpopulation has some additional specific features that make them preferred cargos, since it is becoming increasingly clearer in the field that there is immense heterogeneity among MTs. Post-translational modifications and MT associated proteins, which vastly alter stability and coordination of motor proteins (Hammond et al. 2008, Yu et al. 2015, McKenney et al. 2016, Monroy et al. 2018), might also influence which MTs serve as cargos versus transportation tracks in β cells.

Revised version: Interestingly, blocking kinesin results in a striking accumulation of MT in the cell center where they are normally nucleated at MTOCs, which include the centrosome and the Golgi, in differentiated β cells the latter being the main MTOC. Thus, sliding MTs originate from the MTOC area. At the same time, FIB-SEM analysis did not detect many MTs associated with MTOCs in physiologically normal β cells (Muller et al. 2021). This implies that MTs are typically rapidly dissociated from MTOCs so that they become available for transport by sliding. It is worth mentioning that for long-distance transport by sliding, cargo MTs must be short, otherwise MT buckling and not long-distance transport will occur (Straube et al. 2006). Interestingly, shorter MTs have been observed in high glucose conditions (Muller et al. 2021) when MT are nucleated more actively (Trogden et al. 2019) and transported more frequently (this paper). Possibly, nucleated MTs are detached from MTOCs before they achieve a length that would prevent their transport. An intriguing possibility has been proposed, indicating that in high glucose conditions, MTs might undergo severing by katanin (Muller et al. 2021). This process could generate MT fragments, potentially facilitating their role as cargos with increased ease of transport. It is also possible that sliding MT subpopulation has some additional specific features that make them preferred cargos, since it is becoming increasingly clearer in the field that there is immense heterogeneity among MTs. Post-translational modifications and MT associated proteins, which vastly alter stability and coordination of motor proteins (Hammond et al. 2008, Yu et al. 2015, McKenney et al. 2016, Monroy et al. 2018), might also influence which MTs serve as cargos versus transportation tracks in β cells.

Earlier version: Importantly, we observe a prominent effect of peripheral MT loss only after a long-term kinesin depletion (three-four days). This is consistent with our observation that only a minor subset of MT is being moved within each experimental time frame. We postulate that the absence of a peripheral microtubule array in KIF5B-depleted cells is a consequence of prolonged lack of sliding. We hypothesize that MT sliding must contribute to β-cell specific peripheral MT bungle formation during β cell differentiation. We also found that increasing MT sliding does not yield a properly configured MT array: kinesore-treated cells lack aligned peripheral MTs, consistently with kinesore-induced MT looping reported in other cell types (Randall et al. 2017). This indicates that, similar to other part of β cell physiology, the dose of MT sliding has to be precisely tuned to achieve physiologically relevant architecture. It was shown before that exaggerated kinesin-dependent MT sliding causes MT bundling and buckling into aberrant configuration (Straube et al. 2006). We predict that a fine-tuning regulatory pathway must exist to restrict the number of MT sliding events to the cell needs.

Revised version: Importantly, we observe a prominent effect of peripheral MT loss only after a long-term kinesin depletion (three-four days). This is consistent with our observation that only a minor subset of MT is being moved within each experimental time frame. We postulate that the absence of a peripheral MT array in KIF5B-depleted cells is a consequence of prolonged lack of sliding. We hypothesize that MT sliding must contribute to β-cell specific peripheral MT bungle formation during β-cell differentiation. We also found that increasing MT sliding does not yield a properly configured MT array: kinesore-treated cells lack aligned peripheral MTs, consistently with kinesore-induced MT looping reported in other cell types (Randall et al. 2017). This indicates that, similar to other parts of β-cell physiology, the dose of MT sliding has to be precisely tuned to achieve physiologically relevant architecture. It was shown before that exaggerated kinesin-dependent MT sliding causes MT bundling and buckling into aberrant configuration (Straube et al. 2006). We predict that a fine-tuning regulatory pathway must exist to restrict the number of MT sliding events to the cell needs.

Consistent with this idea, MT sliding is sensitive to metabolic regulation: our data indicate that MT sliding is activated on a short-term basis after glucose stimulation. While the deep understanding of the role on MT sliding in GSIS requires further studies, it is plausible to suggest two potential functions for this process in glucose-stimulated cells. One, given that peripheral MTs are destabilized in high glucose (Ho et al. 2020), we suggest that a long-term function for MT sliding, needed to replace MT population at the cell periphery and restore the pre-stimulus MT organization (Fig. 7B). Because the amount of MT polymer on every glucose stimulation changes only slightly (Zhu et al. 2015, Muller et al. 2021) and we detect MT loss from the periphery only after a prolonged blocking of sliding, we reason that this function could be essential to maintain long-term β-cell fitness and prepare cells for repeated rounds of stimulation. Two, as a potential short-term function at each stimulation round, MT sliding within the peripheral bundle itself could rearrange fragmented MTs within this array suppressing its role in IG withdrawal from specific secretion sites (Fig. 7B). In this scenario, MT sliding would tilt the balance between positive and negative MT-dependent regulation of GSIS toward enhanced secretion at each stimulation.

Added in the revised version: Our finding provides an example of a phenomenon where a slight change in MT configuration is physiologically significant. This is not an exception because subtle MT defects often have dramatic consequences. For example, in mitotic spindles, a tiny overgrowth of MT ends during metaphase, which causes them to attach to both kinetochores rather than just one, is very significant for the efficiency of chromosome segregation, causing aneuploidy and cancer. The changes in β-cell MT networks that we are reporting are much stronger: the effect on the peripheral MT network accumulated over three days of KIF5B depletion is dramatic (Fig 2 B, C). Short-term gross MT network configurations after a single glucose stimulation are harder to detect, but is consistent with previous reports that MTs at the cell periphery are destabilized and fragmented upon high glucose stimulus (Ho et al 2020, Muller et al 2021), and that preventing this MT rearrangement completely blocks GSIS (Zhu et al 2015, Ho et al 2020).

Added in the revised version: One of the most fascinating features of insulin secretion regulation is that the amount of generated insulin granules significantly exceeds the normal physiological needs for insulin secretion (∼100 times more than needed). At the same time, even slightly facilitated glucose depletion can be devastating for the human body. Accordingly, the excessive insulin content of a β cell resulted in the development of multiple levels of control, preventing excessive secretion. Our previous data suggest that the peripheral MT array provides one of those mechanisms. This study indicates that MT sliding is necessary to form the proper peripheral network in the long term. Short-term glucose-induced changes in the peripheral MT array likely need to be subtle to prevent over-secretion. Thus, we are not surprised that a dramatic effect of sliding inhibition is only detectable by our approaches after the changes in the MT network accumulate over time.

Earlier version: On a final note, it is important to evaluate the phenomenon reported here in light of the dual role of KIF5B as IG transporter and MT transporter and the coordination of those two roles in IG transport and availability for secretion. Our results indicate that KIF5B is needed for the formation of peripheral MT bundle which we have shown to restrict secretion (Bracey et al. 2020, Ho et al. 2020). At the same time, it is well established that KIF5B transports IGs and KIF5B loss of function impairs GSIS (Meng et al. 1997, Varadi et al. 2002, Cui et al. 2011). After a prolonged KIF5B inactivation, a loss of peripheral readily-releasable IG should be expected due to two factors: because there is no MT bundle to prevent over-secretion and IG depletion, and because there is no new IGs being transported from the Golgi area. In contrast, physiological activation of kinesin by glucose (Donelan et al. 2002, Varadi et al. 2003) would both promote replenishment of IG through non-directional transport through the cytoplasm and restoration of peripheral MT array to prevent over-secretion.

Revised version: On a final note, it is important to evaluate the phenomenon reported here in light of the dual role of KIF5B as IG transporter and MT transporter and the coordination of those two roles in IG transport and availability for secretion. Our results indicate that KIF5B is needed for the formation of the peripheral MT bundle, which we have shown to restrict secretion (Bracey et al. 2020, Ho et al. 2020). At the same time, it is well established that KIF5B transports IGs and KIF5B loss of function impairs GSIS (Meng et al. 1997, Varadi et al. 2002, Cui et al. 2011). After a prolonged KIF5B inactivation, a loss of peripheral readily-releasable IG should be expected due to two factors: because there is no MT bundle to prevent over-secretion and IG depletion, and because there is no new IGs being transported from the Golgi area. In contrast, physiological activation of kinesin by glucose (Donelan et al. 2002, Varadi et al. 2003) would both promote replenishment of IG through non-directional transport through the cytoplasm and restoration of the peripheral MT array to prevent over-secretion.

In conclusion, here we add another very important cell type to the list of systems that employ KIF5B-dependent MT sliding to build functional cell-type-specific MT networks. This system is unique because in this case MT sliding is metabolically regulated and activated on a single-minute time scale by nutrition triggers.

Materials and Methods

Key reagents

Cell Lines

Earlier version: MIN6 cells between passage 40-60 were utilized (Miyazaki et al. 1990, Ishihara et al. 1993). Cells were maintained in 25 mM glucose Dulbecco’s modified eagle medium (DMEM) (Life Technologies, Frederick, MD) supplemented with 10% fetal bovine serum (FBS), 0.001% β-mercaptoethanol, 0.1 mg/ml penicillin, and 0.1 mg/ml streptomycin in 5% CO2 at 37 degrees C.

Revised version: MIN6 cells between passage 40-60 were utilized (Miyazaki et al. 1990, Ishihara et al. 1993). Cells were maintained in 25 mM glucose Dulbecco’s modified eagle medium (DMEM) (Life Technologies, Frederick, MD) supplemented with 10% fetal bovine serum (FBS), 0.001% β-mercaptoethanol, 0.1 mg/ml penicillin, and 0.1 mg/ml streptomycin in 5% CO2 at 370C.

Mice

Earlier version: Mouse usage followed protocols approved by the Vanderbilt University Institutional Animal Care and Use Committee for G.G. and I.K. Mice were euthanized by isoflurane inhalation. C57BL/6J and Kif5b+/− (Cui et al. 2011) mice were from The Jackson Laboratory (Bar Harbor, ME). Conditional knockout Kif5bfl/−:RIP2-Cre mice were bred and genotyped as previously described (Cui et al. 2011).

Revised version: Mouse usage followed protocols approved by the Vanderbilt University Institutional Animal Care and Use Committee for G.G. and I.K (protocol #M18000195). Mice were euthanized by isoflurane inhalation. C57BL/6J and Kif5b+/− (Cui et al. 2011) mice were from The Jackson Laboratory (Bar Harbor, ME). Conditional knockout Kif5bfl/−:RIP2-Cre mice were bred and genotyped as previously described (Cui et al. 2011).

Islet Isolation and Cell/Islet Culture

Islets were isolated from 8-to 16-week-old mice. Briefly, ∼2 mL of 0.8 mg/mL collagenase P (MilliporeSigma, St. Louis, MI) in Hanks’ balanced salt solution (Corning, Corning, NY) was injected into the pancreas through the common bile duct. The pancreas was digested at 37°C for 20 min. Islets were handpicked into RPMI 1640 media with 11 mmol/L glucose (Gibco, Dublin, Ireland) plus 10% heat-inactivated (HI) FBS (Atlanta Biologicals, Flowery Branch, GA) and cultured at 37°C with 5% CO2. For MIN6 cells, DMEM with 25 mmol/L glucose, 0.071 mmol/ L β-mercaptoethanol (MilliporeSigma), 10% HI FBS, 100 µU/mL penicillin, and 100 µg/mL streptomycin (Gibco) was used.

Reagents and antibodies

Primary antibodies for immunofluorescence were: mouse anti-β-tubulin (Sigma-Aldrich, 1:1000), rabbit anti-β-tubulin (Sigma-Aldrich, 1:1000) and rabbit anti-KIF5B (Abcam), Alexa488-, Alexa568-, and Alexa647-conjugated highly cross-absorbed secondary antibodies (Invitrogen). Coverslips were mounted in Vectashield Mounting Medium (Vector Labaratories). Cells were treated with indicated drugs for three hours unless otherwise indicated. Drugs used were: Kinesore (Tocris Bioscience).

shRNA sequence

The KIF5B-targeting shRNA [shRNA KIF5B] #1, [TL510740B, 5’-ACTCTACGGAACACTATTCAGTGGCTGGA] and [shRNA KIF5B] #2, [TL51074CB 5’ – AGACCGTAAACGCTATCAGCAAGAAGTAG] are in the plasmid backbone pGFP-C-shLenti and were from Origene (Rockville, MD). The non-targeting shRNA control, was pGFP-C-shLenti also from Origene.

DNA Constructs

Cloning

The Scr shRNA -mEmerald-Tubulin, KIF5B shRNA#1-mEmerald-tubulin, KIF5B shRNA#2-mEmerald-tubulin were all generated from their respective TGFP containing constructs. Using the NotI and PmeI sites the TGFP was swapped for mEmerald-Tubulin-C-18.

The FKBP-mCherry-KIF5B(568-964) construct (gift from Kristen Verhey, University of Michigan), has previously been previously described (Ravindran et al. 2017).

By using site directed mutagenesis, we made 8 point mutations in the tail domain to change residues RKRYQ to AAAYA in the ATP independent MT binding domain. The point mutations were sufficient to rescue MT sliding in the cell. As previously published, this disrupts the tail domain to bind to the acidic e hook of the MT tail. Point mutations were introduced using a site directed mutagenesis kit, In-Fusion® Snap Assembly (Takara).

Lentiviral Transduction and Transfection

Lentivirus production and infection followed standard methods (Huang et al. 2018). MIN6 were treated with a given shRNA expressing a mEmerald-tubulin/cytosolic tgfp marker for 96hrs prior to imaging to achieve KD efficiency. For non-viral vectors, MIN6 cells were transfected using Amaxa Nucleofection (Lonza) and experiments were conducted 24 to 48 hours thereafter.

Image Acquisition

Immunofluorescence microscopy of fixed samples

Fixed samples were imaged using a laser scanning confocal microscope Nikon A1r based on a TiE Motorized Inverted Microscope using a 100X lens, NA 1.49, run by NIS Elements C software. Cells were imaged in 0.05µm slices through the whole cell.

Live cell imaging

Cells were cultured on 4-chamber MatTek dishes coated with 10 µg/µl fibronectin and transduced 96hrs or transfected 48 h before experiment. For live-cell imaging of MT sliding, cells were transfected with Emerald-Tubulin and imaged using a Nikon TiE inverted microscope equipped with 488- and 568-nm lasers, a Yokogawa CSU-X1 spinning disk head, a PLAN APO VC 100x NA1.4 oil lens, intermediate magnification 1.5X, and CMOS camera (Photometrics Prime 95B), 405 Burker mini-scanner, all controlled by Nikon Elements software.

Photobleaching Assay

Earlier version: ∼1x10^6 MIN6 cells were transfected with 1µg of mEmerald-tubulin or transduced with lentiviral KIF5B shRNA with mEmerald-tubulin as a reporter and attached to glass dishes coated with fibronectin for up to 96hrs. On the SDC microscope, the ROI tool in NIKON elements was used to place two ROI’s ∼5µm apart at either end of the cell. These regions were assigned to be photobleached with the equipped 405nm mini scanner laser leaving a fluorescent patch over the middle which we termed the “fluorescent belt”. After the regions were photobleached cells were then acquired for 5mins, across 7 optical slices (0.4µm step size) in 10 second interval between frames.

Revised version: ∼1×106 MIN6 cells were transfected with 1µg of mEmerald-tubulin or transduced with lentiviral KIF5B shRNA with mEmerald-tubulin as a reporter and attached to glass dishes coated with fibronectin for up to 96hrs. On the SDC microscope, the ROI tool in NIKON elements was used to place two ROI’s ∼5µm apart at either end of the cell. These regions were assigned to be photobleached with the equipped 405nm mini scanner laser leaving a fluorescent patch over the middle which we termed the “fluorescent belt”. After the regions were photobleached cells were then acquired for 5mins, across 7 optical slices (0.4µm step size) in 10 second interval between frames.

Sun Tag Rigor Kinesin and Tracking of MT sliding

Earlier version: SunTag system for MT lattice fiducial marks was adapted from (Lu et al. 2016). ∼1x10^6 MIN6 cells were co transfected with 1µg of the ScFv-GCN4-HaloTag-GB1-NLS, and 0.5µg of the pcDNA4TO-K560-E236A-24xGCN4 plasmid (K560RigorE236A-SunTag (Tanenbaum et al. 2014)). After 24hrs the cells were washed with 1x PBS and the media replaced with KRB containing 2.8mM glucose for 1 hour, following a second incubation with HALO dye of choice (Promega) for 30mins. Cells were imaged in 1 focal plane for 2 mins with 100ms exposure time and no delay in acquisition.

Revised version: SunTag system for MT lattice fiducial marks was adapted from (Lu et al. 2016). ∼1×106 MIN6 cells were co transfected with 1µg of the ScFv-GCN4-HaloTag-GB1-NLS, and 0.5µg of the pcDNA4TO-K560-E236A-24×GCN4 plasmid (K560RigorE236A-SunTag (Tanenbaum et al. 2014)). After 24hrs the cells were washed with 1× PBS and the media replaced with KRB containing 2.8mM glucose for 1 hour, following a second incubation with HALO dye of choice (Promega) for 30mins. Cells were imaged in 1 focal plane for 2 mins with 100ms exposure time and no delay in acquisition.

Image Processing and presentation

Figure 1: (D-F) Maximum intensity projections of spinning disc confocal stacks through the ventral 1.2µm of the cell for each time point. (G-I) Single focal plane spinning disc confocal time frames.

Figure 1- Supplemental Figure 1: (A-C) Single focal planes from laser scanning confocal stacks.

Figure 2: (A-C) Maximum intensity projections of laser scanning confocal stacks over the ventral 1µm of the cell. The tubulin channel (top) is shown as inverted gray-scale images. The KIF5B channel (bottom) was pseudo-colored cyan.(F-H) Thresholded single focal plane images under the nucleus of the cell selected for directionality analysis.

Figure 2- Supplemental Figure 1: (A, B) Maximum intensity projections of laser scanning confocal stacks over the ventral 1µm of the islet. Inverted gray-scale images. (D, E) Thresholded single focal plane images under the nucleus of the cell selected for directionality analysis.

Figure 2- Supplemental Figure 2. (A-D) Maximum intensity projections of laser scanning confocal stacks through the whole cell. The tubulin channel (top) is shown as inverted gray-scale images. The BFP-motor cannel (bottom) was pseudo-colored cyan.

Figure 2- Supplemental Figure 3. Single-plane laser confocal images shown as inverted gay-scale images (A1-A2), thresholded images (A4, A5), and Matlab analysis outcomes (A8). Images in (A6-A7) are non-to-scale schematics.

Figure 2- Supplemental Figure 4. Single-plane laser confocal images shown as inverted gay-scale images (A1-A2), thresholded images (A3, A4), and Matlab analysis outcomes (A5, A6).

Figure 3: Maximum intensity projections of spinning disc confocal stacks through the ventral 1.2µm of the cell for each time point. (C-G).

Figure 4: Cells shown are maximum intensity projections of laser scanning confocal stacks. through the ventral 1µm of the cell. The tubulin channel (top) is shown as inverted gray-scale images (A-C). Thresholded single focal plane images under the nucleus of the cell selected for directionality analysis (F-H).

Figure 4- Supplemental Figure 1: (A-B) Cells shown are maximum intensity projections of laser scanning confocal stacks. through the ventral 1µm of the cell. The tubulin channel (top) is shown as inverted gray-scale images. mCherry-labeled KIFDN constructs (bottom) are pseudo-colored magenta.

Figure 5: Cells shown are maximum intensity projections of laser scanning confocal stacks. of the ventral 1µm of the cell. The tubulin channel is shown as inverted gray-scale images (A-B). Thresholded single focal plane images under the nucleus of the cell selected for directionality analysis (D-E).

Figure 5 - Supplemental Figure 2: Cells shown are maximum intensity projections of laser scanning confocal stacks. of the bottom 1µm of the cell. The tubulin channel is shown as inverted gray-scale images.

Figure 6: Maximum intensity projections of spinning disc confocal stacks through the ventral 1.2µm of the cell for each time point (A-B). Single focal plane spinning disc confocal time frames (D-E).

For all images, whole-image contrast for each channel was adjusted equally within each figure.

Quantitative Image Analysis

Analysis of MT sliding after photobleaching

Maximum intensity projections of spinning disc confocal stack from the FRAP assay were analyzed. The photobleached cell areas outside the “fluorescent belt” were thresholded to mask out cell background and insignificant amounts of low-signal MTs which presumably represent polymerizing MT ends. The area of high-signal MTs moving into the bleached area from the fluorescent belt at every time pint of the video was quantified as a proxy for MT sliding efficiency.

Sun Tag Rigor Kinesin Tracking Analysis

Earlier version: The acquired image was processed through Imaris Microscopy Image Analysis Software (Oxford Instruments), where the fiducial marks were tracked. Utilizing the spot tracking feature, we filtered the fiducial marks based on intensity and employed the integrated Autoregressive Motion algorithm. Both the Max Distance and Gap sizes were enabled to precisely calibrate the tracks. Subsequently, we refined the tracks by filtering them based on intensity, thereby eliminating noise and particles that intermittently enter and exit focus. We sought to normalize the behaviors by comparing the MT movements over 5-second intervals and calculated the displacement of a given fiducial mark (see MATLAB scripts below). In total over 60,000 tracks were detected, and the segmented displacement of ∼25,000 of those tracks were calculated.

Revised version: The acquired image was processed through Imaris Microscopy Image Analysis Software (Oxford Instruments), where the fiducial marks were tracked. Utilizing the spot tracking feature, we filtered the fiducial marks based on intensity and employed the integrated Autoregressive Motion algorithm. Both the Max Distance and Gap sizes were enabled to precisely calibrate the tracks. Subsequently, we refined the tracks by filtering them based on intensity, thereby eliminating noise and particles that intermittently enter and exit focus. We sought to normalize the behaviors by comparing the MT movements over 5-second intervals and calculated the displacement of a given fiducial mark (see MATLAB scripts below). In total over 60,000 tracks were detected, and the segmented displacement of ∼25,000 of those tracks were calculated. The 5 second displacement was binned at 0.05um intervals and the % of distribution for each bin was calculated for each cell and summarized in histograms (Fig 1 - Supplemental Figure 1B, Figure 6F). Displacements of less than 0.15µm/5sec (below the resolution limit) were considered indicative of stationary MTs (Figure 1J, Figure 6G). Displacements of greater than 0.3µm/5sec were considered significant displacements indicative of motile MTs. (Figure 1K, Figure 6H).

Removed in the revised version: The 5 second displacement was binned at 0.05um intervals and the % of distribution for each bin was calculated for each cell and summarized in histograms (Fig 1 -Supplemental Figure 1B, Figure 6F). Displacements of less than 0.15µm/5sec (below the resolution limit) were considered indicative of stationary MTs (Figure 1J, Figure 6G). Displacements of greater than 0.3µm/5sec were considered significant displacements indicative of motile MTs. (Figure 1K, Figure 6H).

MATLAB Script: Msdanalyzer, Segmentation

The position of all tracked fiducial spots were exported from Imaris to excel. The MSDanalyzer was developed by Nadine Tarantino et al, and adapted by Kai Bracey, Pi’Illani Noguchi, and Alisa Cario (Vanderbilt University) to normalize the tracks in time. Tracks were segmented into displacements over 5s and binned as shown in the results.

MATLAB Script: MT Directionality

Earlier version: Oversampled images were deconvolved using the Richardson and Lucy Deconvolution algorithm. Images were masked and threshold (IsoData) in ImageJ. The MT directionality script was applied in MATLAB. Only the outer 1µm of MTs were taken for binning and quantification purposes.

Revised version: Oversampled images were deconvolved using the Richardson and Lucy Deconvolution algorithm. Images were masked and thresholded (IsoData) in ImageJ. The MT directionality script was applied in MATLAB. Only the outer 1µm of MTs were taken for binning and quantification purposes.

KIF5B depletion quantification

Single focal plane laser scanning confocal images were quantified in ImageJ. Cell outlines were used as masks, and cells positive and negative for GFP expression (shRNA marker) were analyzed. Mean KIF5B immunostaining intensity per cell was measured. To avoid bias of potential staining fluctuation between data points, the intensity of each GFP-positive cell was normalized to the averaged intensity of GFP-negative cells in the same field of view.

Statistics and reproducibility

For all experiments, n per group is as indicated by the figure legend and the scatter dot plots indicate the mean of each group and error bars indicate the standard error of the mean. All graphs and statistical analyses were generated using Excel (Microsoft) and Prism software (Graphpad). Statistical significance for all in vitro and in vivo assays was analyzed using an unpaired t-test, one-way ANOVA with Sidak’s multiple comparisons test, Kolmogorov-Smirnov test as indicated in the figure legends. For each analysis p <0.05 was considered statistically significant, and *p < 0.05, **p < 0.01, ***p< 0.001, ****p<0.0001.